```python
# Initialize Otter
import otter
grader = otter.Notebook("Numpy.ipynb")
```

# Numpy

Similar to lists, NumPy arrays allow us to store data in a sequential and ordered fashion. Here are some key features of NumPy arrays:
1. **Homogeneous Data Types**: In a NumPy array, all elements must be of the same data type. This uniformity allows NumPy to optimize performance and memory usage, whereas Python lists can contain elements of mixed types.

2. **Fixed Size**: Once created, the size of a NumPy array is fixed. If you need to change the size, you typically create a new array. This fixed size helps in optimizing computations and memory management.

3. **Shape and Dimensions**: Each element in a NumPy array has a well-defined shape, which is crucial for performing vectorized operations. For instance, in a 2D array (matrix), all rows have the same number of columns, ensuring consistent dimensionality.

A major advantage of using NumPy arrays is their ability to execute computations more swiftly and efficiently with NumPy’s functions, as opposed to using lists. NumPy is highly valued in scientific computing and is extensively used in data analysis and machine learning. 

```python
import numpy as np
```

## Create a numpy array

Numpy is used to work with arrays. The array in numpy is called **ndarray**. We can create a numpy array by using `array` method. 

```python
my_arr = np.array([1, 3, 5, 10])
my_arr
```

```python
type(my_arr)
```

### Question 1
Create a numpy array to store the values 5, 10, 15 in a variable `arr`. 

```python
arr = ...
arr
```

```python
grader.check("q1")
```

### Question 2

Create a numpy array to store the names of your favorite three fruits. Store your answer in a variable, `my_fruits`. 

```python
my_fruits = ...
my_fruits
```

```python
grader.check("q2")
```

<!-- BEGIN QUESTION -->

### Question 3 

Previously, I noted that a NumPy array should only contain elements of the same data type. However, in the code example below, I created a NumPy array with both integer and string types. Why isn’t NumPy throwing an error in this case? Does the absence of an error imply that NumPy can handle elements of varying data types? What is happening under the hood?

```python
arr = np.array([10, "hello"])
arr
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 0-D arrays (Scalars)

One value stored in a numpy array are 0-D array or scalar. 

```python
arr = np.array([56])
arr
```

Numpy arrays have `shape` property that tells us the number of elements in each dimension. 

```python
arr.shape
```

### 1-D arrays 

An array that has 0-D arrays as its elements  is called 1-D array. 

```python
arr = np.array([5, 6, 7, 8, 9])
arr
```

```python
arr.shape
```

### 2-D arrays

An array that has 1-D arrays as its elements is called 2D array.


For eg: \begin{bmatrix}
10 & 20 & 30 \\
1 & 1 & 1 \\
\end{bmatrix}

```python
mat = np.array([[10, 20, 30], [1, 1, 1]])
mat
```

```python
mat.shape
```

<!-- BEGIN QUESTION -->

### Question 4

What is the cause of error in the following code snippet? 

Hint: Check the `shape` of each element of the numpy array. 

arr = np.array([[10, 20], [30]])

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 5

Create a 2-dimensional array which looks like the following: 

\begin{bmatrix}
10 & 20 \\
11 & 11 \\
100 & 200 \\
\end{bmatrix}

Store your answer in the variable, `your_ans`. What is the shape of the resultant array, `your_ans`? 

```python
your_ans = ...
your_ans
```

```python
grader.check("q5")
```

### Question 6
3-D arrays

An array that has 2-D arrays (matrices) as its elements is called a 3-D array. 

Store the following 3D array in a variable, `arr3d`.  What is the shape of the resultant array? 

<img src="pics/numpyTensors.jpg" style="width:300px" />

```python
arr3d = ...

arr3d
```

```python
arr3d.shape
```

```python
grader.check("q6")
```

### Comparison operators in Python

We've already covered numerical operators like +, -, *, /, //, **, and % in our discussion on expressions. Now, let’s explore another category of operators known as comparison operators. These operators are used to compare two values and produce a Boolean result (either True or False).

| Operator | Meaning | Example | Value | 
| -----   | ---- | ----- | ----- |
| < | less than | 3 < 2 | True |
| <= | less than or equal to | 2 <= 2 | True |
| > | greater than | 3 > 20 | False |
| >= | greater than or equal to | 3 >= 3 | True |
| == | equal to | 50 == 50 | True |
| != | not equal to | 4 != 4 | False |

<!-- BEGIN QUESTION -->

### Question 7

What will be the result of the following expressions? Please determine this without executing any code.

expr = ((5 * 2) == (2 * 4 + 2))
expr

_Type your answer here, replacing this text._

<!-- END QUESTION -->



```python
10e298429 # a very large number is interpreted as infinity in python
```

```python
10e-292393 # a very small number is interpreted as zero in python
```

Python stores floating-point numbers without guaranteeing exact precision.

```python
import math
(math.sqrt(2) ** 2) == 2 # Do not compare between two floating numbers in python; always check if they are close enough
```

```python
np.isclose(math.sqrt(2) ** 2, 2 )
```

### Boolean operators in Python

Boolean expressions are those that evaluate to either True or False.

Python offers boolean operators as well, specifically and and or. These operators work with boolean values. For example, the expression `p and q` utilizes a boolean operator where both p and q are boolean expressions. 

| p | q | p and q | 
| --- | --- | --- |
| True | True | True | 
| True | False | False |
| False | True | False | 
| False | False | False |


| p | q | p or q | 
| --- | --- | --- |
| True | True | True | 
| True | False | True |
| False | True | True | 
| False | False | False |

<!-- BEGIN QUESTION -->

### Question 8

What will be the result of the following expressions? Please determine this without executing any code.

expr = (4 < 5) and (5 == 5)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 9

What will be the result of the following expressions? Please determine this without executing any code.

expr = (5 < 3) or (9 != 9)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Scalar computation on a numpy

When a scalar value (just one number) is added to a numpy array, then that value is added to all the elements of the numpy array at the same time, and produce a new numpy aray as an output, as illustrated in the figure below: 

<img src="pics/broadcast1.jpg" width=500px />

```python
arr = np.array([1, 2, 3])
arr * 2
```

<!-- BEGIN QUESTION -->

### Question 10

What is the output of the following python code?  Please determine this without executing any code.

arr = np.array([1, 2, 3]) + 2
arr

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 11

What is the output of the following python code?  Please determine this without executing any code.

arr = np.array([10, 20, 30]) / 10
arr

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 12

What is the output of the following python code?  Please determine this without executing any code.

arr = np.array([1, 2, 3]) < 3
arr

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 13

What is the output of the following python code?  Please determine this without executing any code.

arr = np.array([1, 2, 3]) == 3
arr

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 14

Given the temperatures of five cities in Celsius—14, 53, 19, and 25—write a Python program to convert these temperatures to Fahrenheit. The formula for converting Celsius to Fahrenheit is:

$temp\_in\_fahr = \frac{9 * temp\_in\_celcius}{5} + 32$

Store your answer in a variable, `temps_in_fahr`. 

Hint: 
1. Create a numpy array, which stores the temperatures in celcius
2. Then use the scalar computation trick to compute the temperatures in fahrenheit all at once

```python
temps_in_cel = ...
temps_in_fahr = ...
temps_in_fahr 
```

```python
grader.check("q14")
```

### Computation between two numpy arrays

Sometimes, we need to perform computations on two arrays, such as adding or multiplying them. These operations are known as elementwise addition or multiplication, where each element in one array is combined with the corresponding element in the other array. To perform these operations, the two arrays must have the same shape.

For example, consider the table below showing the number of students graduating from various colleges at a university:

| College name | Number of students graduating in 2019 | Number of students graduating in 2022 |
| --- | :----: | :---: |
| College of Science and Mathematics | 457 | 532 |
| College of Arts | 385 | 433 | 
| College of Business | 573 | 501 |
| College of Education | 382 | 293 |


Create a Python code to determine the difference in the number of students graduating from each college in 2019 compared to 2022.

```python
studs_in_2019 = np.array([457, 385, 573, 382])
studs_in_2022 = np.array([532, 433, 501, 292])

diff = studs_in_2022 - studs_in_2019
diff 
```

Keep in mind that for two NumPy arrays to be subtracted from one another, they must have the same shape. In this example, both `studs_in_2019` and `studs_in_2022` have a shape of (4,).

### Question 15

Suppose you are given two matrices A and B, where  $A = \begin{bmatrix}
1 & 2 \\
1 & 1 \\
10 & 20 \\
\end{bmatrix}$ and $B = \begin{bmatrix}
1 & 1 \\
2 & 2 \\
3 & 3 \\
\end{bmatrix}$ . Write a python code to add these two matrices together and store your answer in a variable, `C`. 

```python
A = ...
B = ...
C = ...
C
```

```python
grader.check("q15")
```

### Numpy methods 

#### Range (np.arange)

A range is an array of numbers arranged in either increasing or decreasing order, with each number separated by a consistent interval. Ranges are quite useful in various situations, making it important to understand them. You can create ranges using the `np.arange` method, which can accept one, two, or three arguments: the start, end, and step size.

If you provide one argument to `np.arange`, it is treated as the end value, with the start set to 0 and the step to 1 by default. When you provide two arguments, they specify the start and end values, with a default step of 1. Supplying three arguments explicitly sets the start, end, and step values.

```python
np.arange(10) # here start = 0 and step = 1 are set by default
```

```python
np.arange(3, 9) # here start = 3, end = 8, step = 1 by default
```

```python
np.arange(3, 30, 5) # start = 3, end = 29, step = 5 
```

The start, end, and step values can all be positive or negative and can be whole numbers or fractions. For example:

```python
np.arange(1.5, -2, -0.5) # start = 1.5, end = -1.99999.... , step = -0.5
```

### Accessing elements of a numpy array

Like python lists and strings, numpy array also maintain index for each element. The element of a numpy array can be accessed using these indices. For example: 

```python
arr = np.arange(10)
arr
```

```python
arr[0] # returns the first element of the numpy array
```

```python
arr[4] # returns the fifth element of the numpy array
```

You can also do slicing on a numpy array

```python
arr[4: 9]
```

```python
arr[:5]
```

```python
arr[1: : 2]
```

### Question 16

You are given the following matrix, $A = \begin{bmatrix}
1 & 2 & 3 & 4 \\
10 & 20 & 30 & 40 \\
100 & 200 & 300 & 400 \\
-10 & -20 & -30 & -40 \\
-1 & -2 & -3 & -4 \\
\end{bmatrix}$


Your task is to write a python code to only access the elements of 3rd row (100, 200, 300, 400) from the matrix A. Your answer should be of shape (4, )

```python
A = np.array([
             [1, 2, 3, 4],
             [10, 20, 30, 40],
             [100, 200, 300, 400],
             [-10, -20, -30, -40],
             [-1, -2, -3, -4]
            ])

A
```

```python
A.shape
```

```python
your_ans = ...
your_ans
```

```python
grader.check("q16")
```

### Question 17

Given the array A (see Q. 16), write a python code to access first two rows of the matrix. Your answer should be $\begin{bmatrix}
1 & 2 & 3 & 4 \\
10 & 20 & 30 & 40 \\
\end{bmatrix}$ 
of shape (2, 4)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q17")
```

### Question 18

Given the array A (see Q. 16), write a python code to access the last two rows of the matrix. Your answer should be $\begin{bmatrix}
-10 & -20 & -30 & -40 \\
-1 & -2 & -3 & -4 \\
\end{bmatrix}$ 
of shape (2, 4)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q18")
```

### Question 19

Given the array A (see Q. 16), write a python code to access the first two columns of the matrix. Your answer should be $\begin{bmatrix}
1 & 2  \\
10 & 20  \\
100 & 200  \\
-10 & -20  \\
-1 & -2  \\
\end{bmatrix}$
of shape (5, 2)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q19")
```

### Question 20

Given the array A (see Q. 16), write a python code to access the last two columns of the matrix. Your answer should be 
$\begin{bmatrix}
3 & 4 \\
30 & 40 \\
300 & 400 \\
-30 & -40 \\
-3 & -4 \\
\end{bmatrix}$
of shape (5, 2)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q20")
```

### Question 21

Given the array A (see Q. 16), write a python code to access the first and the last rows of the matrix. Your answer should be 
$\begin{bmatrix}
1 & 2 & 3 & 4 \\
-1 & -2 & -3 & -4 \\
\end{bmatrix}$
of shape (2, 4)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q21")
```

### Question 22

Given the array A (see Q. 16), write a python code to access the first and the last columns of the matrix. Your answer should be 
$ \begin{bmatrix}
1  & 4 \\
10   & 40 \\
100  & 400 \\
-10 & -40 \\
-1  & -4 \\
\end{bmatrix}$
of shape (5, 2)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q22")
```

### Question 23

Given the array A (see Q. 16), write a python code to access every second row of the matrix starting from first row i.e, first row, third row, fifth row and so on. Your answer should be 
$\begin{bmatrix}
1 & 2 & 3 & 4 \\
100 & 200 & 300 & 400 \\
-1 & -2 & -3 & -4 \\
\end{bmatrix}$
of shape (3, 5)

```python
your_ans = ...
your_ans
```

```python
your_ans.shape
```

```python
grader.check("q23")
```

### Question 24

Given the array A (see Q. 16), write a python code to access the element 200. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q24")
```

### Question 25

Given the array A (see Q. 16), write a python code to access the element -4. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q25")
```

### Question 26

Given the square matrix $A = \begin{bmatrix}
1 & 2 & 3 \\
40 & 50 & 60  \\
-7 & -8 & -9  \\
\end{bmatrix}$. Write a python code to extract the diagonal elements of the matrix. Your answer should be a numpy array. 

```python
A = ...
your_ans = ...
your_ans
```

```python
grader.check("q26")
```

### Numpy methods

Numpy module provides us many methods that can be applied on numpy array. 

Each of the functions shown below takes an array as an argument and returns a single value. 

```python
arr = np.array([1, 2, 3])
arr
```

```python
np.sum(arr) # add all elements of an array
```

```python
np.mean(arr) # returns the mean of all the elements of the array
```

```python
np.prod(arr) # returns the product of all the elements of the array
```

```python
arr = np.array([0, 5, 3, 0])
np.all(arr) # returns True if all elements are True or non-zero values
```

```python
barr = np.array([True, True, True])
np.all(barr)
```

```python
barr = np.array([True, True, True, False])
np.all(barr)
```

```python
arr = np.array([0, 5, 3, 0])
np.any(arr) # returns True if any of the elements are True or non-zero value 
```

```python
barr = np.array([False, False, False])
np.any(barr)
```

```python
barr = np.array([False, False, False, True])
np.any(barr)
```

```python
arr = np.array([0, 5, 3, 0, 8])
np.count_nonzero(arr) # counts the number of True or non-zero elements in the array
```

Each of the following functions takes an array as an argument and returns a numpy array

```python
arr = 3 ** np.arange(5)
arr
```

```python
np.diff(arr) # returns the difference of the consecutive elements 
```

```python
arr = np.array([0.5, 0.239, 0.8238])
np.round(arr) # returns the round of each number to the nearest integer
```

```python
arr = np.array([1, 2, 3, 4])
np.cumprod(arr) # a cumulative product
```

```python
np.cumsum(arr) # a cumulative sum
```

```python
np.exp(arr) # exponentiate each element 
```

```python
np.log(arr) # natural logarithm of each element
```

```python
np.sqrt(arr) # square root of each element
```

```python
arr = np.arange(10, 1, -1)
arr
```

```python
np.sort(arr) # sort elements in ascending order
```

### Boolean Arrays and Boolean Masking



#### Boolean array: 
A numpy array where all values of the array are Boolean (True or False). 

Examples of boolean array: 

```python
barr = np.array([True, False, True, True, False])
barr
```

```python
type(barr), type(barr[0])
```

```python
arr = np.array([[5, 6], [9, 10]])
arr
```

```python
barr = arr > 6
barr
```

#### Methods on boolean array

```python
barr1 = np.array([True, True, False, False, True])
barr2 = np.array([True, False, False, True, True])
```

```python
np.logical_and(barr1, barr2) # elementwise and operator
```

```python
np.logical_or(barr1, barr2) # elementwise or operator
```

### Boolean Masking

A Boolean mask is a technique used to select elements from a NumPy array by applying a Boolean array of the same shape. The resulting array will be one-dimensional and contain only the elements corresponding to the True values in the Boolean array.

For example: 

<img src="pics/bm1.drawio.png" width=700 />

```python
arr = np.array([10, 20, 30, 40])
barr = np.array([True, True, False, True])
arr[barr]
```

Example: 

<img src="pics/bm2.drawio.png" width=700 />

```python
arr = np.array([[10, -20, 30], [-1, 3, 5], [-6, -4, 9]])
arr
```

```python
arr > 0
```

```python
arr[arr > 0]
```

### Question 27

Write a python code to find the sum of all positive integers in the matrix $A= \begin{bmatrix}
1 & -2 & -3 & 4 \\
1 & 1 & -1 & -1 \\
2 & -2 & -1 & 0 \\
\end{bmatrix}$

```python
A = np.array([[1, -2, -3, 4], [1, 1, -1, -1], [2, -2, -1, 0]]) 
your_ans = ...
your_ans
```

```python
grader.check("q27")
```

### Question 28

Write a python code to find the sum of the positive integers present in the second row of the matrix A (defined in Question 26). 

```python
your_ans = ...
your_ans
```

```python
grader.check("q28")
```

### Question 29

Write a python code to count the total number of zeros present in the matrix $A= \begin{bmatrix}
1 & 5 & 0 & 4 \\
1 & 3 & 0 & 0 \\
2 & -2 & -1 & 0 \\
4 & 0 & 9  & 0  \\
2 & 9 & 0 & 0 \\
\end{bmatrix}$

```python
A = ...
A
```

```python
your_ans = np.sum(A == 0)
your_ans
```

```python
grader.check("q29")
```

### Question 30

Write a python code to replace all the negative numbers with 0 in the given array $A= \begin{bmatrix}
1 & -4 & -1 & 4 \\
1 & 3 & -8 & 1 \\
\end{bmatrix}$. 

Your resultant array A will be $\begin{bmatrix}
1 & 0 & 0 & 4 \\
1 & 3 & 0 & 1 \\
\end{bmatrix}$. 

```python
A = np.array([[1, -4, -1, 4], [1, 3, -8, 1]])
A
```

```python
mask = ...
A[mask] = ...
A
```

```python
grader.check("q30")
```

### Miscelleneous Questions

### Question 31

#### Estimate $\pi$ using Leibniz's formula

Leibniz has given us formula to estimate the value of $\pi$ using the following equation: 
$$\pi = 4 * \left(1 - \frac{1}{3} +  \frac{1}{5} -  \frac{1}{7} + \frac{1}{9} -  \frac{1}{11} + ... \right)$$

We are going to verify if this equation is true by computing the right hand side of equation, if the result is approximately 3.14 (close to the value of $\pi$) then the equation is correct. 

Instead of finding the sum infinitely, we will caculate it upto large value (say 100001), then the equation becomes:
$$\left(1 - \frac{1}{3} + \frac{1}{5} - \frac{1}{7} + \frac{1}{9} - \frac{1}{11} + ... - \frac{1}{99999} + \frac{1}{100001} \right)$$

Now we will separate this equation by grouping the negative fractions and positive fractions together respectively. So the equation becomes:
$$\left(1 + \frac{1}{5} + \frac{1}{9} + ... + \frac{1}{100001} \right) - \left(\frac{1}{3} + \frac{1}{7} + \frac{1}{11} + ... + \frac{1}{99999} \right)$$

Therefore, let's first calculate the positive fractions first. 
$$\left(1 + \frac{1}{5} + \frac{1}{9} + ... + \frac{1}{100001} \right)$$

We see pattern in the denominator 1, 5, 9, ..., 100001. 
Therefore, we can store these denominators in a numpy array. 

Then, we can divide them by 1, and then add them together. That will be the output for the sum of all positive fractions. 

```python
positive_fractions = ...
left_equation = ...
```

```python
negative_fractions = ...
right_equation = ...
```

```python
pi = ...
pi
```

```python
grader.check("q31")
```

### Question 32

#### Estimate $\pi$ using Walli's Formula

About half a centure before Leibniz, John Wallis gave an expression to estimate the value of $\pi$ as an infinite product. 

$$\pi = 2 \cdot \left(\frac{2}{1} \cdot \frac{2}{3} \cdot \frac{4}{3} \cdot \frac{4}{5} \cdot \frac{6}{5} \cdot \frac{6}{7} \cdots \right) $$

This is a product of "even/odd" fractions. Let's use arrays to multiple a million of them, and see if the product is close to $\pi$. 

$$\pi = 2 \cdot \left(\frac{2}{1} \cdot \frac{4}{3} \cdot \frac{6}{5} \cdots \frac{100000}{99999} \right) \cdot \left(\frac{2}{3} \cdot \frac{4}{5} \cdot \frac{6}{7} \cdots \frac{100000}{100001} \right) $$

We start by creating an array of even numbers 2, 4, 6, and so on upto 100000. Then we create two lists of odd numbers: 1, 3, 5, 7, ... upto 99999 and 3, 5, 7, ... upto 100001


```python

left_equation = ...

right_equation = ...

pi = 2 * left_equation * right_equation
pi
```

```python
grader.check("q32")
```

### Question 33

A numpy array of integers, named `nvidia` contains the rough adjusted closing monthly stock prices (in dollars) of NVIDIA corporation, a key innovator in computer graphics, and AI technology, for the year 2023. It has total 12 items.


```python
nvidia = np.array([19.52, 23.20, 27.76, 27.74, 37.82, 42.28, 46.71, 49.34, 43.48, 40.77, 46.76, 49.51])
nvidia
```

Write a Python code to answer each of the following questions.

#### Question 33(a)
What is the stock price value in October? Round your answer upto two decimal points.

```python
your_ans = ...
your_ans
```

```python
grader.check("q33")
```

### Question 33(b)
The stock price change for each consecutive month from Jan to Dec. For eg: the change of stock price between month of Feb and Jan is 23.2 - 19.52 = 3.68. Round your answer upto two decimal points. Your final answer should be a numpy array with 11 elements. Store your answer in the variable, `nvidia_diff`. 

```python
nvidia_diff = ...
nvidia_diff
```

```python
grader.check("q33b")
```

### Question 33(c)

Check if the stock prices ever rose by more than 4 dollars between consecutive months. Your solution should include Python code that returns either True or False (without explicitly writing True or False in the answer).

```python
nvidia_diff
```

```python
your_ans = ...
your_ans
```

```python
grader.check("q33c")
```

### Question 33(d)

Determine how many times the stock prices increased by more than 4 dollars between consecutive months. Your answer should include Python code that returns either True or False (without directly writing True or False in your answer).

```python
your_ans = ...
your_ans
```

### Question 33(e)

Determine how many times the stock prices ever **rose** by less than 4 dollars between consecutive months. Your solution should include Python code that returns either True or False (without explicitly writing True or False in the answer).

```python
nvidia_diff
```

```python
your_ans = ...
your_ans
```

```python
grader.check("q33e")
```

<!-- BEGIN QUESTION -->

### Question 34:

For each of the Python expressions below, write the output when the expression is evaluated. If the
expression evaluates to an array, you should format your answer like so: `array([..., ..., ...])`. If the expression fails to run, then write `Error` as your answer. You
may assume the standard import:

```import numpy as np```  

Do not run any code to answer these questions. 

#### Question 34(a) 

```np.array([5, 6, -3]) > 4```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(b)

```np.sum(np.array([5, 2, 0]) < 3)```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(c)

```np.array([1, 2, 3]) - np.array([1, 1, 1])```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(d)

```np.array([2, 2, 2]) - 2 ```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(e)

```np.array([4, 4, 5]) - np.array([2, 2])```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(f)

```np.arange(1, 10, 2)```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(g)

```
arr = np.arange(10, 100, 10)
new_arr = arr[np.arange(1, 5)]
print(new_arr)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(h)

```np.arange(1, 10, 3)```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(i)

```
arr = np.arange(10, 100, 10)
new_arr = arr[np.arange(1, 7, 2)]
print(new_arr)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34j

```
arr = np.arange(10, 100, 10)
new_arr = arr[np.arange(1, 10, 2)]
print(new_arr)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34k

```
np.cumsum(np.array([1, 2, 4]))
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34l

```
arr = np.array([3, 5, 7])
arr[np.array([True, False, False])]
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34m

```
arr = np.array([[1, 1], [-2, 3]])
arr[arr > 0]
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(n)

```
arr = np.array([[1, 1], [-2, 3]])
arr[arr < 0] = 1000
arr
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 34(o)

```
"csci" + str(round(1461.78))
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->



Tests for q1 (points: None)
```
>>> isinstance(arr, np.ndarray)
True
```


Tests for q14 (points: None)
```
>>> temps_in_fahr[-1] == 77
np.True_
```
```
>>> len(temps_in_fahr) == 4
True
```


Tests for q15 (points: None)
```
>>> C.shape == (3, 2)
True
```
```
>>> C[-1, -1].item() == 23
np.True_
```


Tests for q16 (points: None)
```
>>> your_ans.shape == (4,)
True
```
```
>>> your_ans[-1].item() == 400
True
```


Tests for q17 (points: None)
```
>>> your_ans.shape == (2, 4)
True
```
```
>>> your_ans[-1, -1].item() == 40
True
```


Tests for q18 (points: None)
```
>>> your_ans.shape == (2, 4)
True
```
```
>>> your_ans[-1, -1].item() == -4
True
```


Tests for q19 (points: None)
```
>>> your_ans.shape == (5, 2)
True
```
```
>>> your_ans[-1, -1].item() == -2
True
```


Tests for q2 (points: None)
```
>>> isinstance(my_fruits, np.ndarray)
True
```
```
>>> len(my_fruits) == 3
True
```


Tests for q20 (points: None)
```
>>> your_ans.shape == (5, 2)
True
```
```
>>> your_ans[-1, -1].item() == -4
True
```


Tests for q21 (points: None)
```
>>> your_ans.shape == (2, 4)
True
```
```
>>> your_ans[-1, -1].item() == -4
True
```


Tests for q22 (points: None)
```
>>> your_ans.shape == (5, 2)
True
```
```
>>> your_ans[-1, -1].item() == -4
True
```


Tests for q23 (points: None)
```
>>> your_ans.shape == (3, 4)
True
```
```
>>> your_ans[-1, -1].item() == -4
True
```


Tests for q24 (points: None)
```
>>> your_ans.item() == 200
True
```


Tests for q25 (points: None)
```
>>> your_ans.item() == -4
True
```


Tests for q26 (points: None)
```
>>> your_ans[-1].item() == -9
True
```
```
>>> isinstance(your_ans, np.ndarray)
True
```


Tests for q27 (points: None)
```
>>> your_ans.item() == 9
True
```


Tests for q28 (points: None)
```
>>> your_ans.item() == 2
True
```


Tests for q29 (points: None)
```
>>> your_ans.item() == 8
True
```


Tests for q30 (points: None)
```
>>> A.shape == (2, 4)
True
```
```
>>> A[0, 2].item() == 0 and A[0, 1].item() == 0 and (A[1, 2].item() == 0)
True
```


Tests for q31 (points: None)
```
>>> np.isclose(3.14, np.round(pi, 2)).item()
True
```


Tests for q32 (points: None)
```
>>> np.isclose(3.14, np.round(pi, 2)).item()
True
```


Tests for q33 (points: None)
```
>>> np.isclose(40.77, your_ans).item()
True
```


Tests for q33b (points: None)
```
>>> len(nvidia_diff) == 11
True
```
```
>>> np.isclose(nvidia_diff[-1], 2.75).item()
np.True_
```


Tests for q33c (points: None)
```
>>> your_ans.item() == True
True
```


Tests for q33e (points: None)
```
>>> your_ans == 3
True
```


Tests for q5 (points: None)
```
>>> your_ans.shape == (3, 2)
True
```


Tests for q6 (points: None)
```
>>> arr3d.shape == (3, 3, 3)
True
```
