```python
# Initialize Otter
import otter
grader = otter.Notebook("HW3.ipynb")
```

Due Date: 18 Apr 2025 by 11:59 pm

Grading Rubric:

| Question | Points |
| ----- |---- |
| 1(a) | 5 |
| 1(b) | 5 |
| 1(c) | 10 |
| 1(d) | 10 |
| 2(a) | 10 |
| 2(b) | 10 |
| 2(c) | 10 |
| 2(d) | 10 |
| 2(e) | 10 |
| 3(a) | 5 |
| 3(b) | 15 |
| Total | 100 |

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd

from matplotlib.ticker import PercentFormatter
```

Instructions: 

Submission: Please upload the notebook (.ipynb) file on D2L --> HW -- > HW03 --> HW03_Dropbox

1. The HW assignment is individual-based assignment. However, you are encouraged to discuss with your classmates, but DO NOT show your code and copy from others. It will be considered as plagiarism, and proper actions will be taken as per the syllabus.

2. Just because your question has passed the autograder test, doesn't mean necessarily that your answer is correct. So make sure you answer those questions carefully. The instructor has some hidden tests, and only when your question passes those tests successfully during grading, you will be given full points for the question.

3. If you are stuck somewhere, feel free to reach out to me. 

Objectives: To practice and learn about the following topics: 
1. Histogram
2. Intro to Probability

# 1. Histogram

The data below shows the median household income (in tens of thousands) in different cities of LA(louisiana), TX(texas) and NM (new mexico) for 900 families (300 each). 

```python
income = pd.read_csv('data/income.csv')
income.head()
```

```python
income.shape
```

```python
# Custom formatter to remove the percentage sign
def custom_formatter(x, pos):
    return f'{x * 100:.1f}' 
```

```python
fig = plt.figure(figsize=(20,5))
bins = np.arange(0, 151, 10)
import matplotlib.ticker as mtick

plt.subplot(1, 3, 1)
plt.hist(income.query('State == "NM"')['Median Income'], bins = bins, edgecolor = 'black', density = True)
plt.xlabel("Income of families (dollars)")
plt.ylabel("% of Income per dollar")
plt.title("New Mexico")
plt.gca().yaxis.set_major_formatter(mtick.FuncFormatter(custom_formatter))
plt.grid()

plt.subplot(1, 3, 2)
plt.hist(income.query('State == "TX"')['Median Income'], bins = bins, edgecolor = 'black', density = True)
plt.xlabel("Income of families (dollars)")
plt.ylabel("% of Income per dollar")
plt.title("Texas")
plt.gca().yaxis.set_major_formatter(mtick.FuncFormatter(custom_formatter))
plt.grid()

plt.subplot(1, 3, 3)
plt.hist(income.query('State == "LA"')['Median Income'], bins = bins, edgecolor = 'black', density = True)
plt.xlabel("Income of families (dollars)")
plt.ylabel("% of Income per dollar")
plt.title("Louisiana")
plt.gca().yaxis.set_major_formatter(mtick.FuncFormatter(custom_formatter))
plt.grid()

plt.show()
```

The three histograms plotted above depict the distributions of median incomes for families in the three states. All bars are 10k dollars wide, and all bar heights are represented as income per dollar. 

Answer the following questions solely on the histogam provided above. You should not write any python code for this. 

<!-- BEGIN QUESTION -->

### Question 1(a) [5 points]

Find the percentage of families in the state of Louisiana who have income at least 50k but less than 70k. (Do not write any code for this, consider the height of bar [60, 70) to be 0.9)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(b) [5 points]

By looking at the graph above, is it possible to accurately determine the percentage of families in the state of New Mexico who have income atleast 20k but less than 45k? Why or why not? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(c) [10 points]

Find the **number** of families that have income of atleast 20k in the state of Louisiana? (Consider the height for the bar [10, 20) is 0.45)

If your answer is in decimals, then consider only integer part as your answer. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(d) [10 points]

Assume there are no families with income 110k in louisiana, and there are no families with income 100k in texas. 

Fill in the blanks below with the smallest range that can be determined from the histograms shown above. 

The percentage of families in Texas that have income higher than all the families in Louisiana is between ______ and _______. 

Consider the height of the bar [100, 100), [110, 120), [120, 130) and [130, 140] for the state of Texas is 0.25, 0.2, 0.15, 0.15 respectively 

Explain your answer. 

Note: This is difficult question. Please make sure you discuss with your friends (do not show the solution) or reach out to me if you are stuck. Do not skip it please. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 2. Probability 

SFA dining food hall provides multiple stations serving a wide variety of menu items from homestyle favorites to diverse international flavors. Emily, a student at SFA, likes to get 3 servings of food items on her plate from any of those stations, and chooses each station independently and at random with replacement. For example, she could randomly get food items from salad station from three times, and thus end up with three servings on her plate, or, she may choose to randomly get a food item from Homestyle station, another from Innovate station, and third from Sweet Shoppe. 

The following table lists the food stations, and the probability that Emily chooses each. 

| Station | Probability |
| --- | --- |
| Sweet Shoppe | 0.35 |
| Homestyle | 0.05 |
| Deli | 0.15 |
| Innovate | 0.1 |
| Salad | 0.15 | 
| Pasta | 0.2 |

Note: you may leave your answers as mathematical expressions. **Show your work**

<!-- BEGIN QUESTION -->

### Question 2(a) [10 points]

What is the probability that Emily ends up picking items in this exact order: 

Homestyle, Salad, Sweet Shoppe?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2(b) [10 points]

What is the probability that Emily does not get Pasta in any of her three servings? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2(c) [10 points]

What is the probability that Emily gets either three servings from Innovate station or three servings from Deli station? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2(d) [10 points]

What is the probability that Emily ends up with two servings from Homestyle station and one from Sweet Shoppe? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2(e) [10 points]

What is the probability that Emily gets one or more food items from Salad station? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### 3. Probability (using simulation)

Isaac Newton was consulted about the following problem by Samuel Pepys, who wanted the information for the gambling purposes. You will act as Newton and try to answer the following questions. Write the simulation code to derive the result. 

<!-- BEGIN QUESTION -->

### Question 3(a) [5 points]

Find the probability that atleast two 6's appear when twelve fair dice are rolled. 

```python
dice_sides = np.arange(1, 7)
```

```python
def one_expt():
    """
    Expt: Randomly pick six numbers from dice_sides
    returns: count of the number of 6 appeared after performing the expt
    """
    ...
```

```python
repetitions = 50000
total_outcomes = []

for ... in ... :
    ... .append(...)

```

```python
total_outcomes = np.array(total_outcomes)
```

```python
count_no_6_appeared = ...
count_only_one_6_appeared = ...

prob = 1 - count_no_6_appeared / repetitions - count_only_one_6_appeared / repetitions
prob
```

```python
grader.check("q3a")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

## Question 3(b) [15 points]

### Monty Hall Problem 

Monty Hall is trying out a new version of his game, with rules as follows. The contestant gets to choose one of four doors. One door has a car behind it, another has an
apple, another has a book, and another has a goat. All 24 (4!) permutations for which door
has which prize are equally likely. In order from least preferred to most preferred, the
contestant’s preferences are: goat, apple, book, car.
Monty, who knows which prize is behind each door, will open a door (other than the
contestant’s initial choice) and then let the contestant choose whether to switch to
another unopened door. Monty will reveal the least preferred prize (among the 3 doors
other than the contestant’s initial choice) with probability p=0.7, the intermediately preferred
prize with probability (1 − p) = 0.3, and the most preferred prize never.

The contestant decides in advance to use the following strategy: Initially choose door
After Monty opens a door, switch to one of the other two unopened doors, randomly
choosing between them (with probability 1/2 each).

Write a simulation code to find the probability that the contestant will get the car.

Hint: 
1. Let the contestant pick a door randomly with equal probability.
2. Sort the remaining doors, such that the last two items will be the least desired item by the contestant, and the items that Monty would pick with probability 0.3 and 0.7 respectively.
3. Pick randomly an item from the last two items of the sorted remaining doors. This will be the Monty pick. 
4. Pick randomly from the left-over two items for contestant pick after he/she switches
5. Check if the contestant pick is a car door or not and return accordingly.

```python
# Lets' assume that car is denoted by the digit 1, book by 2, apple by 3 and goat by 4 in ascending order.
doors = {'car': 1, 'book': 2, 'apple': 3, 'goat': 4}
```

```python
behind_closed_doors = np.array([1, 2, 3, 4])
np.random.shuffle(behind_closed_doors) # shuffle the items
behind_closed_doors
```

```python
def one_expt():
    np.random.shuffle(behind_closed_doors) # Randomly shuffling the items behind the doors that only Monty knows about

    #contestant_pick is the item that the contestant will pick randomly. Use np.random.choice to select randomly an item from behind closed doors
    contestant_pick = np.random.choice(behind_closed_doors)

    doors_left_to_pick = ...
    doors_left_to_pick = np.sort(doors_left_to_pick)

    #monty_pick is the item that the Monty will pick randomly. Use np.random.choice to select randomly an item from behind closed doors.
    monty_pick = np.random.choice(..., p=...)

     
    doors_for_contestant = ...

    contestant_pick_after_switch = ...
    
    if(contestant_pick_after_switch == 1):
        ...
    else:
        ...
```

```python
repetitions = 50000
total_outcomes  = []

for x in np.arange(repetitions):
    ...
```

```python
total_outcomes = np.array(total_outcomes)
```

```python
prob = ...
prob
```

```python
grader.check("q3b")
```

<!-- END QUESTION -->



Tests for q3a (points: None)
```
>>> (prob > 0.5 and prob < 0.8).item() == True
True
```


Tests for q3b (points: None)
```
>>> (0.3 < prob < 0.4).item() == True
True
```
