```python
# Initialize Otter
import otter
grader = otter.Notebook("DataStructures.ipynb")
```

# Data Structures

In everyday life, we design specific methods for storing data to make it more accessible and manageable. For example, plates are often arranged sequentially in the kitchen for easy access, customers might be lined up in a queue for movie tickets, and leaves and fruits are organized on a tree by branches and stems. Similarly, in computer programming, we use various data structures to organize and store data in a way that allows us to access and manipulate it efficiently, reflecting how we structure data in the real world.

Python offers several pre-defined, built-in data structures that are useful for data analysis. In this lesson, we will explore some of these structures, including strings, lists, tuples, and dictionaries.

## 1. Strings

Strings are the zero or more sequence of characters in python. They are surrounded by double quotation marks (or single quotation marks). 
For eg:

```python
print("Hello World")
print("") # this is empty string which has 0 character
```

### Indexing in String

Each character in a string maintains an index starting from 0 (shown in the table below):

|Index --> | 0 | 1 | 2 | 3 | 4 |
|----| ---- | --- | --- | ---- | ---- |
| | H | e | l | l | o |

We can use indexing to access any character of a string. 

For eg: To access the first character of the string, we can do: 

```python
my_str = "Hello"
my_str[0] # get character at position index 0
```

### Question 1:

Consider the string `Python Programming`, write a python code to extract the character `h`. Store your answer in the variable `your_ans`.

```python
my_str = "Python Programming"
your_ans = ...
print(your_ans)
```

```python
grader.check("q1")
```

### Question 2:

Use the same variable, `my_str` in Question 1, and extract the character `a` from the string stored in the variable. Store your answer in the variable `your_ans`.

```python
your_ans = ...
print(your_ans)
```

```python
grader.check("q2")
```

### String length

To get the length of a string, we can use the in-built python function, `len`

```python
len(my_str) # Hello contains total 5 characters
```

### Question 3:

Consider the string, `I love Python`. Write a python code to extract the character `n` from the string. Store your result in the variable, `your_ans`. 
Can you solve the question by using the `len` method to extract the index of last character of the string? 

```python
my_str = "I love Python"
your_ans = my_str[len(...) - ...]
print(your_ans)
```

### Question 4:

Write a python code to find the total number of characters in a string, `Be the change that you wish to see in the world.` Store your result in a variable `your_ans`. 

```python
my_str = "Be the change that you wish to see in the world."
your_ans = ...
print(your_ans)
```

```python
grader.check("q4")
```

### Check String

To check if a substring is present in a string, we can use `in` keyword. 

```python
'o' in my_str
```

```python
my_name = "Dipak Singh"
'Sing' in my_name
```

### Question 5:

Write a python code to check whether the word, `crow` exists in the sentence, `Why did the scarecrow win an award? Because he was outstanding in his field!`. Store your answer in the variable, `your_ans`. Your answer should store boolean value, True or False. 

```python
funny_str = "Why did the scarecrow win an award? Because he was outstanding in his field!"
your_ans = ...
print(your_ans)
```

```python
grader.check("q5")
```

### Question 6:

Write a python code to check whether the word, `dog` exists in the variable, `funny_str` you defined in Question 5. Store your answer in the variable, `your_ans`. 
Hint: Use `not in` keyword

```python
your_ans = ...
print(your_ans)
```

```python
grader.check("q6")
```

### Slicing String

To extract a substring from a string, you can use slicing. Slicing involves specifying the starting and ending index positions within brackets, where the ending index is exclusive. For example, consider the following sentence:
|Index --> | 0 | 1 | 2 | 3 | 4 | 5 | 6 | 7 | 8 | 9 | 10 | 11 | 12 |  13 |14 | 15 | 16 | 17 | 18 | 19 | 20 |
| --- |--- | ---| ---|--- |--- | ---| ---| ---| ---| ---| ---| ---| ---| ---| ---| ---| ---|--- |--- |--- |--- |
| | I | | l | o | v | e |  | D | a | t | a | | A | n | a | l | y | t | i | c | s |

```python
my_var = "I love Data Analytics"
my_var[2:6] # the index number 6 is exclusive here
```

```python
my_var[:5] # By leaving out the start index, the range will start at the first character
```

```python
my_var[12:] # By leaving out the last index, the range will end at the last character
```

```python
# You can even use negative indexing where the index starts from the last character with index position -1 

my_var[-1]
```

```python
my_var[-6]
```

You can even slice characters of a string by adding a step size to it. For eg: 

```python
my_var[::2]
```

Consider the string, 'dog'. What do you think, will be the output of the following python code? Explain why. 

```python
my_str = "dog"
print(my_str[-1::-1])
```

### Question 7:

Consider the string, `supercalifragilisticexpialidocious`. Write a python code to extract the substring `fragilistic`. Store your answer in the variable, `your_ans`.

```python
my_str = "supercalifragilisticexpialidocious"
your_ans = ...
print(your_ans)
```

```python
grader.check("q7")
```

### Question 8

Use the variable `my_str` defined in Question 7, and extract the substring, `super`. Store your answer in the variable, `your_ans`.

```python
your_ans = ...
print(your_ans)
```

### Question 9:
Imagine you have a string that represents a fictional secret code: code = "abracadabra". The secret message is hidden in a special pattern within this code.

Your task:

- Extract the secret message from the code string using slicing, where the message is made up of every 3rd character.
- Reverse the extracted secret message to reveal the final hidden message.

 Store your answer in the variable, `your_ans`. 

```python
code = "abracadabra"
secret_msg = ...
hidden_msg = ...
print(hidden_msg)
```

```python
grader.check("q9")
```

### String Methods

Python provides various in-built functions to be applied on strings. A complete list of such functions can be found [here](https://docs.python.org/3/library/stdtypes.html#string-methods)

#### Change string to upper/lower case

The `upper` method is used to convert a string to upper case. 

```python
my_var = "Hello World"
my_var.upper() # Notice how the method is invoked by using the string followed by a dot.
```

```python
my_var.lower() # similarly, you can use lower to change to lower case
```

### Count String

The `count` method is used to count the number of times a substring appears in a string. 

```python
tongue_twister = "I scream, you scream, we all scream for ice cream."
tongue_twister.count('cream')
```

### Question 10:

Consider the following tweet: 

>Can’t decide if I want coffee or a nap. So, I’m just going to have a Coffee and then take a nap. ☕😴☕😴 #CoFFeeAndNap #Decisions #LivingTheDream #TooTired

Write a Python code to determine the number of times the word coffee appears in a tweet, ensuring that the match is case-insensitive. This means that variations like Coffee, CoFFee, or CoffEE should all be counted. Store your answer in the variable `your_ans`. 

```python
my_tweet = """Can’t decide if I want coffee or a nap. 
So, I’m just going to have a Coffee and then take a nap.
☕😴☕😴 #CoFFeeAndNap #Decisions #LivingTheDream #TooTired"""

your_ans = ...
print(your_ans)
```

```python
grader.check("q10")
```

#### Remove whitespaces from beginning or end of a string

The `strip` method can be used to remove whitespaces at the start or end of a string.

```python
my_var = "    Hello         World            "
my_var.strip() # Notice how it only removes the spaces at the start and the end, but not in the middle
```

```python
my_var = "        Hello World!!!!!!!!!"
my_var.strip('! ') # You can specify multiple characters as a string argument to remove them from the beginning and end of the string.
```

### Question 11:

Write a python code to extract the word `data` by stripping the unwanted characters from the string, `,,,!!!!mldata***&&`. Store your result in the variable, `your_ans`.

```python
my_str = ",,,!!!mldata***&&&"
your_ans = ...
print(your_ans)
```

```python
grader.check("q11")
```

#### Replace String

The `replace` method replaces a string with another string

```python
my_var = "I love CSCI1462"
my_var.replace("CSCI1462", "Data Analytics")
```

### Question 12

Given the string `The company will be hosting a seminar about the new technologies. The technologies are revolutionary and will change the industry`, write a Python code to automatically replace occurrences of the word `technologies` with `innovations` and `seminar` with `workshop`. Store your result in the variable `your_ans`. 

Manual replacement of these words is impractical, especially if the string contains a large number of occurrences, such as 1 million instances. Hence, it's essential to use Python's `replace()` method to handle these replacements efficiently and accurately.

```python
my_str = "The company will be hosting a seminar about the new technologies. The technologies are revolutionary and will change the industry."
your_ans = ...
print(your_ans)
```

```python
your_ans == "The company will be hosting a workshop about the new innovations. The innovations are revolutionary and will change the industry."
```

#### Split String

We can use the `split` method to divide a string into multiple substrings based on a delimiter. This method returns a list containing the resulting substrings.

```python
fruits = "Apple, Banana, Orange, Kiwi"
fruits.split(", ") # Splitting strings based on , character. Note the space after the comma
```

### Question 13

You are given a string that contains a list of comma-separated names. Your task is to split the string into individual names using the `split` method. Store your answer in a variable, `your_ans`. 

```python
list_of_names = 'Bob;Charlie;Victoria;Anson'
your_ans = ...
print(your_ans)
```

```python
grader.check("q13")
```

#### String Concatenation

You can use the `+` operator to combine two strings together. 

```python
s1 = "I love"
s2 = "Data Analytics"
print(s1 + " " + s2)
```

### Question 14
You are building a simple user registration system and need to generate a welcome message for new users. Given the user's first name and last name, your task is to:
1. Create a full name by concatenating the first name and last name with a space in between.
2. Generate a welcome message in the format: 
`Welcome, [Full Name]! We're excited to have you join us.`
3. Print the welcome message.

For eg: 
If the user’s first name is "John" and last name is "Doe", the output should be:

`Welcome, John Doe! We're excited to have you join us.`

```python
first_name = "John"
last_name = "Doe"
full_name = ...
welcome_msg = ...
print(welcome_msg)
```

You can achieve the same result using string formatting, which often makes the code cleaner and more readable. Here’s how you could do it using f-strings.

```python
print(f"Welcome, {full_name}! We're excited to have you join us.")
```

```python
grader.check("q14")
```

#### Escape Character within a String

If you want to add a special character within a string, then you may add them by using an escape character. 

For eg: Suppose we want to print the string, `He said, "How are you?"`

my_var = "He said, "How are you?"" 
print(my_var)

Oops! It seems the Python interpreter is having trouble distinguishing the double quotation marks within the string from the ones that delimit the string itself. To resolve this, you should use the escape character, a backslash (\), to indicate that the double quotation marks inside the string are meant to be part of the string rather than delimiters. Here’s how you can do it:

```python
my_var = "He said, \"How are you?\""
print(my_var)
```

### Question 15

Write a python code to display the following statements: 

```
 Success is not final
 Failure is not fatal
 It is the "courage" to continue that counts
                 - Winston S. Churchill
```

Ensure the two tab spaces precede the attribution line (- Winston S. Churchill). 

Hints:
1. To print messages in Python , utilize the f-string formatting.
2. To include tab spaces, use `\t` as an argument within the print method. Similarly, use `\n` to introduce new line. For instance,

```python
print("Hello\tWorld") # this will add a tab space in between Hello and World
```

```python
print("Hello\nWorld") # this will add a new line in between Hello and World
```

```python
...
```

# 2. Lists

In Python, lists are used to store items in a sequential order. These items can be of various types and are created using square brackets.

```python
fruits = ["Apple", "Banana", "Orange", "Kiwi"]
type(fruits)
```

```python
students = [["Kevin", 23, "484 Bourborn St."], ["Tyler", 34, "95 Esplanade Ave."]]
```

### Question 16:

Write a python code to create a list to store five of your friend's name. Store your list in the variable, `friends`. 

```python
friends = ...
friends
```

```python
grader.check("q16")
```

#### Indexing in Lists

Since lists are ordered, they maintain an index similar to strings. You can access items in a list using these indices.

```python
fruits[0]
```

```python
fruits[:3]
```

```python
students[-1]
```

### Question 17:

Write a python code to print the names of last three friends in your `friends` list you created in Question 16. Store your answer in the variable, `three_friends`. 

```python
three_friends = ...
three_friends
```

```python
grader.check("q17")
```

### Question 18:

Please solve the question 17; using **negative indexing**. 

```python
three_friends = ...
three_friends
```

```python
grader.check("q18")
```

### Question 19:

Write a python code to extract the string `abracadabra` from the list `my_list`. Store your answer in the variable `your_ans`. 

```python
my_list = [56, [192, [395, 539, "abracadabra", 5]], [38, 59, 56]]
your_ans = ...
your_ans
```

```python
grader.check("q19")
```

### Check items in a list

You can check if an item is in a list, using `in` keyword

```python
"Kiwi" in fruits
```

```python
"Mango" in fruits
```

### Question 20:

 Check if the name `Alex` is in your friends list. Store your answer in the variable `your_ans`. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q20")
```

### Length of a list

To determine the number of items in a list, you can use `len` in-built function

```python
len(students)
```

### Changing items in a list

Lists are mutable, meaning you can modify, add, or remove items. To change an item in a list, you can use indexing to assign a new value.

```python
fruits[0]= "Mango"
fruits
```

To insert item in a list, you may use `append` function on a list. 

```python
fruits.append("Peach")
```

```python
fruits
```

### Adding items to a list

To insert item at a particular index of a list, you can use `insert` function. 

```python
fruits.insert(0, "Apple")
fruits
```

To add another list to a list, use `extend` function.

```python
fruits.extend(["Pomegranate", "Pineapple"])
fruits
```

You may also use `+` to add/concatenate two lists together

```python
list1 = ["A", "B", "C"]
list2 = [1, 2, 3]
my_list = list1 + list2
my_list
```

### Question 21:

 Write a python code to add two more friends name `Jeremy`, `Dipak` in your list `friends`.

```python
...
friends
```

```python
grader.check("q21")
```

### Remove items from a list

To remove an item from a list, you may use `remove` function

```python
fruits.remove("Peach")
```

```python
fruits
```

To remove an item based on an index, you may use `pop` function

```python
fruits.pop(1)
fruits
```

```python
fruits.pop() # if no index is given, it will remove the last item from the list
fruits
```

### Question 22:
 Write a python code to remove the second friend from your `friends` list. 

```python
...
friends
```

```python
grader.check("q22")
```

### Question 23:
 Write a python code to remove `Jeremy` and `Dipak` from your `friends` list. 

```python
...
...
friends
```

```python
grader.check("q23")
```

### Copying list

When a list is copied to another list, both lists actually refer to the same items. Therefore, changes made to one list will be reflected in the other.

For eg: 

```python
list1 = [1, 2, 3]
list2 = list1 # here list1 is copied to list2

list2[0] = 100
list1 # although the first item of list2 is changed, however they both refer to the same items, it will reflect in list1
```

If you want to make a separate copy of the list, use `copy` method. 

```python
list1 = [1, 2, 3]
list2 = list1.copy()
list2[0] = 100
list1 # now changing list2 doesn't effect list1
```

```python
list2
```

## 3. Tuples

Like lists, tuples are used to store items in a sequential order. However, tuples are defined using parentheses and are immutable, meaning their contents cannot be changed after creation

```python
colors = ("Green", "Black", "Blue", "Red")
colors
```

```python
type(colors)
```

### Question 24:

Write a python code to create a tuple to store three courses you are taking this semester. Store your answer in the variable `your_courses`. 

```python
your_courses = ...
your_courses
```

```python
grader.check("q24")
```

### Indexing in tuples

Similar to lists, you can use index in tuple to access items in a tuple. 

```python
colors[1]
```

```python
colors[:3]
```

```python
colors[-3]
```

Since tuples are immutable and their values are ordered, you can "unpack" each item of a tuple into individual variables as shown below:

```python
name = ("Dipak", "Singh")
first_name, second_name = name #Unpacking the two values of a tuple into two variables. 

print(f"First item: {first_name} ;  Second item: {second_name}")
```

### Question 25:

Write a python code to unpack the three courses from the variable `your_courses` into variables `course1`, `course2` and `course3`. 

```python
course1, course2, course3 = ...
```

```python
grader.check("q25")
```

# 4. Dictionaries



Dictionaries are used to store items in a special format of `key: value` pair. Dictionaries are changeable, and ordered (since python 3.7+). 

```python
my_info = {"name": "Steve", "age" : 30, "state" : "Texas"}
type(my_info)
```

```python
student = {"name": "Steve", "age" : 30, "courses" : ["CSCI1394", "CSCI924"]}
```

Avoid using items with duplicate keys. Keys are immutable, such as strings can be used as keys in dictionary. 

```python
my_info  = {"name": "Steve", "age" : 30, "state" : "Texas", "name": "Kevin"}
```

```python
my_info
```

### Question 26:

Write a python code to create a dictionary to represent a library catalog where each book is identified by its ISBN number. Each entry in the dictionary should include the book’s title, author, and publication year. Store your result in the variable `library_catalog`

Your dictionary should look like the following:

```
    'ISBN1': {
        'title': 'Sherlock Holmes: A Study in Scarlet',
        'author': 'Sir Arthur Conan Doyle',
        'publication_year': 1887
    },
    'ISBN2': {
        'title': 'Pride and Prejudice',
        'author': 'Jane Austen',
        'publication_year': 1813
    }
```

```python
library_catalog = ...
```

```python
grader.check("q26")
```

### Accessing items of a dictionary

You can access item of a dictionary by refering to it's key name. 

```python
student["name"]
```

```python
student["courses"]
```

You can get all the keys of a dictionary using `keys` function

```python
student.keys()
```

You can get all values of a dictionary using `values` function

```python
student.values()
```

You can get items of a dictionary using `items` function

```python
student.items()
```

You can also access items in a nested dictionary like the following:

```python
nested_dict = {
    'person': {
        'name': 'John Doe',
        'age': 30,
        'address': {
            'street': '123 Main St',
            'city': 'Anytown',
            'zipcode': '12345'
        },
        'contacts': {
            'email': 'johndoe@example.com',
            'phone': '555-1234'
        }
    },
    'job': {
        'title': 'Software Engineer',
        'department': 'Engineering',
        'years_experience': 8
    }
}
```

```python
nested_dict["job"]["title"] # Accessing through multiple keys in dictionary
```

### Question 27:

Write a python code to access the author name with `ISBN2` catalog number from `library_catalog` dictionary. Store your result in variable `your_ans`

```python
your_ans = ...
your_ans
```

```python
grader.check("q27")
```

### Check if an item exists

To check if a key exists in a dictionary, you can use the `in` keyword

```python
"name" in student
```

### Changing items of a dictionary

To change any value of a dictionary, we can refer to it's key name.

```python
student["name"] = "Tyler" # changing the name from Steve to Tyler
student
```

### Question 28:

Change the name of the authors to ` A C Doyle` and `J Austen` respectively to the dictionary `library_catalog`


```python
library_catalog["ISBN1"]["author"] = ...
library_catalog["ISBN2"]["author"] = ...
library_catalog
```

```python
grader.check("q28")
```

### Adding items to a dictionary

You can add an item to a dictionary like the following:

```python
student["country"] = "USA"
student
```

### Question 29:

Add a key `price` to each of the books such that the book Sherlock Holmes costs `$15` and Pride and Prejudice costs `$10` respectively.  

```python
library_catalog["ISBN1"]["price"] = ...
library_catalog["ISBN2"]["price"] = ...
library_catalog
```

```python
grader.check("q29")
```

### Removing items from a dictionary

You can remove item from a dictionary by refering to the key in the `pop` function

```python
student.pop("country")
student
```

### Question 30:

Write a python code to remove the item `price` from the dictionary. 

Hint: First access the inner dictionary and then use `pop` function to remove the item

```python
...
...
library_catalog
```

```python
grader.check("q30")
```

Tests for q1 (points: None)
```
>>> your_ans == 'h'
True
```


Tests for q10 (points: None)
```
>>> your_ans == 3
True
```


Tests for q11 (points: None)
```
>>> your_ans == 'data'
True
```


Tests for q13 (points: None)
```
>>> your_ans[0] == 'Bob'
True
```
```
>>> len(your_ans) == 4
True
```


Tests for q14 (points: None)
```
>>> full_name == 'John Doe'
True
```
```
>>> welcome_msg == "Welcome, John Doe! We're excited to have you join us."
True
```


Tests for q16 (points: None)
```
>>> len(friends) == 5
True
```


Tests for q17 (points: None)
```
>>> len(three_friends) == 3
True
```
```
>>> three_friends[-1] == 'Victor' and three_friends[0] == 'Cynthia'
True
```


Tests for q18 (points: None)
```
>>> three_friends[-1] == 'Victor' and three_friends[0] == 'Cynthia'
True
```
```
>>> len(three_friends) == 3
True
```


Tests for q19 (points: None)
```
>>> your_ans == 'abracadabra'
True
```


Tests for q2 (points: None)
```
>>> your_ans == 'a'
True
```


Tests for q20 (points: None)
```
>>> isinstance(your_ans, bool)
True
```


Tests for q21 (points: None)
```
>>> len(friends) == 7
True
```


Tests for q22 (points: None)
```
>>> len(friends) == 6
True
```


Tests for q23 (points: None)
```
>>> len(friends) == 4
True
```


Tests for q24 (points: None)
```
>>> isinstance(your_courses, tuple)
True
```
```
>>> len(your_courses) == 3
True
```


Tests for q25 (points: None)
```
>>> isinstance(course1, str) and isinstance(course2, str) and isinstance(course3, str)
True
```


Tests for q26 (points: None)
```
>>> isinstance(library_catalog, dict)
True
```
```
>>> 'ISBN2' in library_catalog
True
```
```
>>> 'title' in library_catalog['ISBN1']
True
```


Tests for q27 (points: None)
```
>>> your_ans == 'Jane Austen'
True
```


Tests for q28 (points: None)
```
>>> library_catalog['ISBN1']['author'] == 'A C Doyle'
True
```
```
>>> library_catalog['ISBN2']['author'] == 'J Austen'
True
```


Tests for q29 (points: None)
```
>>> library_catalog['ISBN1']['price'] == '$15'
True
```
```
>>> library_catalog['ISBN2']['price'] == '$10'
True
```


Tests for q30 (points: None)
```
>>> ('price' in library_catalog['ISBN1']) == False
True
```
```
>>> ('price' in library_catalog['ISBN2']) == False
True
```


Tests for q4 (points: None)
```
>>> your_ans == 48
True
```


Tests for q5 (points: None)
```
>>> your_ans == True
False
```


Tests for q6 (points: None)
```
>>> your_ans == True
True
```


Tests for q7 (points: None)
```
>>> your_ans == 'fragilistic'
True
```


Tests for q9 (points: None)
```
>>> hidden_msg == 'bar'
True
```
