```python
# Initialize Otter
import otter
grader = otter.Notebook("HW1.ipynb")
```

```python
import pandas as pd
import numpy as np
```

**Due data: 10 Feb 2025 by 11:59 pm**

Rubric: 
| Question | Points | 
| ---- | --- |
| 1.1 | 10 |
| 1.2 | 10 |
| 1.3 | 10 |
| 2.1 | 10 |
| 2.2 | 10 |
| 3.1 | 10 |
| 3.2 | 10 |
| 3.3 | 10 |
| 4.1 | 10 |
| 4.2 | 10 |
| Total | 100 |

Instructions: 

1. Please upload your notebook (.ipynb) file to D2L under HW >> HW1 Dropbox.
2. The homework assignment is meant to be completed individually. Do not share or copy code with others, as violations will result in consequences outlined in the syllabus.
3. Passing the autograder test does not guarantee that your answer is correct. Be sure to answer questions carefully, as the instructor has additional hidden tests. Full points will only be awarded if your answers pass these tests during grading.
4. If you encounter any difficulties, don't hesitate to contact me for assistance.

Objectives: To practice and learn about the following topics:
1. Variables and Expressions
2. Functions
3. Observation and Causality
4. Data Structures

# 1. Variables and Expressions

<!-- BEGIN QUESTION -->

### Question 1.1 [10 points]

Executing the given code (below) will result in Python generating an error message. 

Identify the issue with the code provided below. Give explanation behind the error. 

10 = 2 * 5

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1.2 [10 points]

For the sequence of code provided, indicate the values of the variables after each statement is executed (i.e., although this is one program, we are executing each statement in separate cells to observe the current values of the variables). If a variable does not have a value at that point, write "N/A"

The first example has been completed for you: 

| Line number | Code | Value of x | Value of y | Value of z |
| ----- | ----- | :----: | :-----: | :----: | 
| 1) |  x = 1 | 1 | N/A | N/A |
| 2) | y = 1 | | | |
| 3) | z = 3 | | | |
| 4) | x = z | | | |
| 5) | z = 5 | | | |
| 6) | x = y + 2 | | | |
| 7) | y = y - 3 | | | |
| 8) | z = x + y | | | |

Note: To verify your solution, execute each line of code in separate cells and use `print(f"x: {x}, y: {y}, z: {z}")` to observe the values of the variables after executing each line.

Format your answer as follows: `1) x: 1, y: N/A, z: N/A`

Here, `1)` indicates the line number, `x: 1` shows the value of x after line 1, `y: N/A` indicates that y has no value yet after line 1, and `z: N/A` shows that z also has no value after line 1.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1.3 [10 points]

Give the code snippet: 

```
x = 10.0
y = (x < 100.5)
```
What is the value of y? Just provide the answer without writing any code.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 2. Functions

### Question 2.1 [10 points]

In Python, there are many string functions available, but most programmers don’t memorize them all. Instead, they usually consult online resources for documentation and examples. Websites like StackOverflow are helpful for finding advice on how to use these methods effectively. For this task, you need to learn about the startswith() function for strings by researching it online.

Here’s the task:

Create a function called `checkFirstName` that takes two inputs: a full name and a search keyword. The function should return `True` if the first name matches the search keyword, and `False` otherwise. Do not use if-else statements, as these haven’t been covered yet. Instead, use the `startswith` method, which returns `True` or `False` based on whether the first name starts with the search keyword. Remember that both the name and the search keyword might be in any case, such as `PeTer gene HerNandez` or `peter gene HERNANDEZ`, where both cases have the first name `Peter`. Your function should work for any full name and search keyword. 

```python
def checkFirstName(...):
    isFirstName = ...
    return isFirstName
```

```python
isFirstNameFound = checkFirstName("Peter Gene Hernandez", "peTEr")
isFirstNameFound
```

```python
grader.check("q2.1")
```

<!-- BEGIN QUESTION -->

### Question 2.2 [10 points]

Strings and numbers are distinct data types. For instance, 5 is a number (integer type), while "5" is a string (string type).

Attempting to add a number (int) and a string (str), such as 5 + "5", results in an error because the operation is not valid.

Fortunately, Python provides built-in functions to convert between these data types: you can convert a number to a string and vice versa.

|Function name | Description | 
| --- | --- | 
| int | convert a string of digits or a float (real number) to an integer (int) value | 
| float | convert a string of digits (perhaps with decimal points) or an int to a float (real number) value | 
| str | convert any number (integer or real number) to a string | 

What is the value of `example` after executing the following code?

```
expr = 2 + int("10") + float("8") # When a float and int is added, the result is float.
example = "Salary " + str(expr)
```

Please provide the answer directly in the markdown cell below, without running the code.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 3. Data Structures

### Question 3.1 [10 points]

Write a Python code to add 600 after 500 in the list provided below. Use the `insert` method to add 600, rather than manually inserting it.

Your final list, `my_list`, should appear as follows:

[1, 2, [30, 40, [500, 600, 700], 40, 30], 2, 1]

```python
my_list = [1, 2, [30, 40, [500, 700], 40, 30], 2, 1]
...
my_list
```

```python
grader.check("q3.1")
```

### Question 3.2 [10 points]

In the 2010s, Cambridge Analytica, a data analytics firm, collected data from millions of Facebook users through an application **without their consent** and aimed to use this information for electoral manipulation. Data from social media can reveal extensive details about an individual's lifestyle, beliefs, and personal information, potentially influencing their political views.

As data analysts, it is important to be cautious of such unethical uses of data.

In this task, you will explore how to extract sensitive information from social media data, though in a ethical manner, since the tweet used in this question is publicly available.

You have a sample of tweets from Donald Trump during his presidency (2017-2021). Your goal is to identify and extract the Twitter user screen name mentioned in his tweets.

Your code should return the string, `henrymcmaster`.

Hint: Use the `user_mentions` key to access the required information.

```python
trump_tweet ={"created_at":["Tue Jun 26 01:35:03 +0000 2018"],"id":[1.01142255594771e+18],"id_str":["1011422555947712513"],"full_text":["Thank you South Carolina. Now let’s get out tomorrow and VOTE for @HenryMcMaster! https://t.co/5xlz0wfMfu"],"display_text_range":[[0],[105]],"entities":{"hashtags":[],"symbols":[],"user_mentions":[{"screen_name":["henrymcmaster"],"name":["Henry McMaster"],"id":[16335288],"id_str":["16335288"],"indices":[[66],[80]]}],"urls":[{"url":["https://t.co/5xlz0wfMfu"],"expanded_url":["https://www.pscp.tv/w/bgGOtTFvTlFsTFJub1dwUXd8MXlvSk1WZHJWQm54Uf-J8fPu1RO4E84ax-LuK1bAbiCpnzBBZmdPfI9FAhGV?t=11s"],"display_url":["pscp.tv/w/bgGOtTFvTlFs…"],"indices":[[82],[105]]}]},"source":[r"<a href=\"http://twitter.com/download/iphone\" rel=\"nofollow\">Twitter for iPhone<\/a>"], "lang":["en"]}
trump_tweet
```

```python
# Do not simply type the name in the blanks. Write a python code to extract that information. 
trump_mentions = ...
trump_mentions
```

```python
grader.check("q3.2")
```

### Question 3.3 [10 points]

Utilize the `replace` string method to convert the text `HitchHiker` into `MatchMaker`. Note that you should replace only specific substrings within the original string, rather than replacing the entire string.

Store your result in the variable, `your_ans`. 

Do NOT replace the entire string with another string. In other words, do not replace entire `HitchHiker` string with `MatchMaker` string. Simply replace substring within `HitchHiker` to make it `MatchMaker`. 

```python
my_str = "HitchHiker"
your_ans = ...
print(your_ans)
```

```python
grader.check("q3.3")
```

# 4. Observation and Causality

<!-- BEGIN QUESTION -->

### Question 4.1 [10 points]

We have learned that `an association between two variables does not necessarily imply causation` (one causes the other). Which of the following examples best illustrates this concept?
Select only one answer. Explain why you selected or rejected each option. Even if your answer is correct, you will not receive credit unless your explanation is accurate.

A) Jon Snow found that customers of the S&V water company had a higher mortality rate from cholera compared to customers of the Lambeth company. This was because Lambeth sourced their water from a cleaner upstream source than S&V, not because dirty water directly caused cholera.

B) Data indicates a link between smoking and lung cancer. This is because smoking directly causes lung cancer.

C) A study by the American Heart Association found a positive correlation between literacy and health (lower literacy, poorer health). However, this is not because literacy itself affects health; rather, it’s due to the fact that schools in impoverished areas are less effective, leading to higher illiteracy rates. Poverty increases health risks because people in poorer conditions can't afford good medical care. Therefore, illiteracy often correlates with poorer health due to poverty.

D) A study shows that literacy is positively associated with health. This is because individuals with poor health are unable to study effectively and therefore are often illiterate.


_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 4.2 [10 points]

To establish the causality between literacy and health relationship, which of the following would be the best way to test? Select only one answer.

A. Run a randomized controlled experiment

B. Brainstorm some confounding factors, and check if any of them have association with literacy and health. 

C. Ask people in the poverty striken neighborhood to find how many of them are illiterate and unhealthy, and if the number is greater, then we can conclude the causality. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->



Tests for q2.1 (points: None)
```
>>> checkFirstName('peter gene hernandez', 'PETER') == True
True
```
```
>>> checkFirstName('petter gene hernandez', 'PETER') == False
True
```
```
>>> checkFirstName('adam sullivan', 'AdAm') == True
True
```


Tests for q3.1 (points: None)
```
>>> my_list[2][2][1] == 600
True
```


Tests for q3.2 (points: None)
```
>>> trump_mentions == 'henrymcmaster'
True
```


Tests for q3.3 (points: None)
```
>>> your_ans == 'MatchMaker'
True
```
