```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab01.ipynb")
```

```python
import numpy as np
import math
```

Grading Rubric:

| Question | Points |
| ----- |---- |
| 1.1 | 10 |
| 1.2 | 10 |
| 2.1 | 10 |
| 2.2 | 10 |
| 2.3 | 10 |
| 2.4 | 10 |
| 3.1 | 10 |
| 3.2 | 10 |
| 4.1 | 10 |
| 4.2 | 10 |
| Total | 100 |

# Welcome to Lab01.

Objectives:
1. Learn print statements
2. Errors in Python
3. Python variables and expressions
4. Python functions

# 1. Print Statements

### Question 1.1 (10 points)

In class, we learned how to show output on the screen using the `print` function. For example, to display the message **I love Data Analytics**, you would use the following `print` statement:

```python
print("I love Data Analytics")
```

In the same way, you can use the `print` statement to show the value of a variable, as demonstrated below:

```python
my_var = 10
print(my_var)
```

In this exercise, we'll explore how to display both strings and variables together on the screen. Often, rather than just printing the value of a variable, we might want to format that value within a string. For example, take a look at the following cell:

```python
value_of_pi = round(math.pi, 2) # stores the value of pi rounded upto two decimal points
print(value_of_pi) 
```

If we want to format the value of pi with the message **The value of pi is 3.14**, we can use Python's [formatted string literal](https://docs.python.org/3/tutorial/inputoutput.html#tut-f-strings) for this purpose.  

Formatted string literals (also called f-strings for short) let you include the value of Python expressions inside a string by prefixing the string with f or F and writing expressions as `{expression}`

For eg: 

```python
print(f"The value of pi is {value_of_pi}")
```

Notice that the value_of_pi variable is enclosed in curly braces within the string so that it can be replaced by the value it holds when using the print statement. 

#### Task: 

Mark is employed at a Data Analytics firm with a yearly salary of $80,000 and receives a 5% bonus for 2023. To determine Mark's total gross income for the end of the year, write a Python code to compute the total amount, including the bonus. Also, display the gross income in the specified format:

`Mark's total gross income is <total_gross_income> dollars in 2023.`

Note: Ensure that you replace `<total_gross_income>` with the variable representing the calculated income. Avoid inserting the amount directly; instead, use the appropriate variable in the print statement.

Your answer should look like the following: 

`Mark's total gross income is 84000.0 dollars in 2023.` 

```python
mark_salary = ...
total_income = ...

print(...)
```

### Question 1.2 (10 points)
Strings consist of zero or more characters, including letters, numbers, symbols, and more. While keyboards typically support a specific set of characters known as ASCII, modern textual data often requires characters from various languages or even emojis, which are not part of the standard ASCII set. To handle these additional characters, programming languages use Unicode, which provides a comprehensive range of characters from across the globe.

For instance, the Unicode representation for the Spanish letter Ñ is U+00D1. To display this character in Python, you can use the following code: print('\U000000D1'). Note that the backslash (\) precedes the argument in the print function, and the plus sign is replaced with zeros, ensuring a total of eight characters after \U.

```python
print('\U000000D1')
```

As a data scientist, working with textual data often involves handling Unicode characters like emojis. For example, examining emojis in movie reviews can provide insights into the reviewer's sentiment. Therefore, it's essential to have a way to read or display these emoji characters.

#### Task: 
Write a Python code to display your favorite emoji that conveys your feelings about taking this course.

Hint: Look at this [website](https://unicode.org/emoji/charts-14.0/full-emoji-list.html) to retrieve the **unicode** for your emoji.

```python
...
```

# 2. Errors in Python

The Python language, akin to human languages, is structured by rules, and transgressions come with consequences. When these rules are violated, errors are generated, unlike in natural languages where flexibility exists.

Making errors is a normal part of programming, even for seasoned developers. Identifying and rectifying errors involves careful analysis of your code. Occasionally, seeking assistance from friends, online platforms like Stack Overflow, GitHub issues, etc., becomes necessary to resolve coding errors.

In this particular scenario, errors have been introduced into my Python code, and I require your assistance to rectify them. Create a new cell and fix them.

### Question 2.1 (10 points)

**ModuleNotFound Error**

```python
import mth
```

Review the final line of the error message above. This message indicates the issue within the code: it states that the module 'mth",' which you are attempting to import, cannot be found. Resolve this error by correcting the module name to 'pandas,' and then attempt to run the code again. If you've successfully addressed the issue, no errors should be encountered.

### Question 2.2 (10 points)

```python
print("hello there
```

This type of error is referred to as a syntax error because there has been a violation of the syntax or structure of the Python code. By carefully reading the error message, you will notice that there is a missing structure at the end of the file (EOF), specifically at the end of this line. The caret (^) symbol also indicates where the structure is broken. Correct this line and rerun the cell. This time, you should not encounter any errors.

### Question 2.3 (10 points)

```python
x + 5
```

### Question 2.4 (10 points)

Kevin, who is new to Python programming, is attempting to create a function called displayTime. This function should take a total number of seconds as input, convert this time into minutes and remaining seconds, and print a message in the format:

`<seconds> seconds is equivalent to <minutes> minutes and <seconds> remaining seconds.`

For eg: `125 seconds is equivalent to 2 minutes and 5 seconds`.


However, Kevin's function contains **syntax and logical errors** and isn't producing the correct output. Assist Kevin in fixing the function to achieve the desired result. 

```python
deaf displayTime(seconds):
    mins = seconds / mins
    remaining_secs = seconds // 60
    print("125 seconds is equivalent to mins and remaining_secs")
```

```python
displayTime()
```

# 3. Expressions

### Question 3.1 (10 points)

Write Python expression in the next cell that evaluates the following expression and save the output in a variable named _result_:
$$2 * (5\frac{5}{2}) + 2^{0.5 * 22} - \frac{7}{45} + 1$$

Please note by $a\frac{b}{c}$,  I mean $a + \frac{b}{c}$ and not $a * \frac{b}{c}$

```python
result = ...
print(result)
```

```python
grader.check("q3.1")
```

<!-- BEGIN QUESTION -->

### Question 3.2 (10 points)

Python adheres to specific rules when naming variables.The rules for naming variables in Python are as follows:

1. Python variable names can consist only of letters (upper-case or lower-case), the underscore symbol _, and numbers.
2. The initial character in a variable name cannot be a number.
3. Variable names cannot contain spaces.
4. Variable names cannot be Python keywords.

Additionally, it is considered good practice to make variable names informative. For example, if you need to store a salary in a variable, it is more meaningful to use a variable named `salary = 60000` than `x = 60000`.

Choose the appropriate variable names from the options below:

    A. 23alpha
    B. sales_tax_2022
    C. number of students
    D. salaryIn$
    E. __totalPopulation
    
There is no need to modify/fix the variable names.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 4. Functions

### Question 4.1 (10 points)

Newton's law of universal gravitation states that every particle attracts every other particle in the universe with a force (in Newton) that is given as follows:
$$F = \frac{G * m_1 * m_2}{d^2}$$
Where G is the gravitational constant, $m_1$ and $m_2$ are the mass of the two objects, and d is the distance between them. 

Write a function which takes mass of any object on earth's surface as input and returns the gravitational force with which the moon attracts that objects towards it according to the formula above. 

Given $G = 6.674 * 10^{-11} Nm^2{kg}^2$, 
mass of moon $ = 6 * 10^{24} kg$ and 
distance between moon and earth's surface is $3.84 * 10^8 m$

```python
... computeForce(...):
    G = ...
    mass_of_moon = ...
    distance = ...
    F = ...
    return ...
```

Now, proceed to execute the function you previously created to determine the force (in Newton) exerted by the moon on my friend, John who is 160 pounds. Store your answer in the variable `force_on_John`. 

Hint: 1 pound is approximately equal to 0.453 kilograms.

```python
force_on_John = ...
```

```python
grader.check("q4.1")
```

### Question 4.2 (10 points)

Write a function named `greet_user` that takes the following parameters:

- name (string) – the name of the user to greet.
- greeting (string, default value: "Hello") – the greeting message.
- punctuation (string, default value: "!") – the punctuation mark to end the greeting.
    
The function should return a greeting string that combines these inputs in the following format:

```greeting, name punctuation```

For example:

greet_user("Alice") should return `"Hello, Alice!"`

greet_user("Bob", greeting="Hi") should return `"Hi, Bob!"`

greet_user("Charlie", punctuation="!!!") should return `"Hello, Charlie!!!"`

greet_user("Dana", greeting="Hey", punctuation=".") should return `"Hey, Dana."`

```python
def greet_user(...):
    ...
```

```python
greet_user("Alice")
```

```python
greet_user("Bob", greeting="Hi")
```

```python
greet_user("Charlie", punctuation="!!!")
```

```python
greet_user("Dana", greeting="Hey", punctuation=".")
```

Tests for q3.1 (points: None)
```
>>> math.isclose(result, 2063.844, abs_tol=0.1)
True
```


Tests for q4.1 (points: None)
```
>>> math.isclose(force_on_John, 0.2, abs_tol=0.1)
True
```
