```python
# Initialize Otter
import otter
grader = otter.Notebook("Expressions.ipynb")
```

```python
import math
```

# Python as a Calculator

We can use python to act as a calculator. Python supports different numerical operators to perform such mathematical operations, such as: 

| Operator | Meaning | Example | Value | 
| -----   | ---- | ----- | ----- |
| * | Multiplication | 3 * 2 | 6 |
| / | Division | 3 / 2 | 1.5 |
| + | Addition | 3 + 2 | 5 |
| - | Subtraction | 3 - 2 | 1 |
| ** | Exponentiation | 3 ** 2 | 9 |
| % | Modulo | 3 % 2 | 1 |

All of the above operators are also called _binary operators_, as they need two data(binary) on each side of the operator (called operands) to compute the result or value. 

When a mathematical operation is made up of either an operator, or variable, or operand and produce an output, we call such an operation as `expression`. For eg: 

```python
5 + 2 * 4
```

```python
x = 5
x * 3
```

### Question 1

Write a python expression to calculate the value of $4^5$. Store your result in a variable `your_ans`. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q1")
```

### Question 2

What is the output of the python expression `17 % 4`? Do not write code for this, simply store your answer in the variable `your_ans` below.

### Question 3

Write a python expression to store the result of the expression $\frac{5 * 6 + 2^5 - 20}{5} $ in a variable `your_ans`. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q3")
```

### Question 4
Write a python code to convert 6 feet to meters, and store your result in a variable `your_ans`. One foot is 0.305 meter. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q4")
```

### Question 5
Write a python code to convert 50 pounds into kilograms, and store your result in a variable `your_ans`. One pound is 0.454 kilogram. 

```python
your_ans = ...
your_ans
```

```python
grader.check("q5")
```

### Question 6

Mark works for a Data Analytics company and have an annual salary of 88,000 USD. He also gets an annual bonus of 10% on top of his salary. What is the total gross annual income of Mark? Store your answer in a variable `mark_income` 

```python
mark_income = ...
mark_income
```

```python
grader.check("q6")
```

### Question 7

Write a python program that converts a given minutes into the number of years and remaining days for the minutes. For simplicity, assume that a year has 365 days. 

For eg: 1000000000 (1 billion) minutes is approximately 1902 years and 214 days. 

```python
mins = 1000000000 
MINS_IN_A_DAY = 24 * 60
DAYS_IN_A_YEAR = 365
total_days = ...
years = ...
remaining_days = ...

print("{} minutes is approximately {} years and {} days.".format(mins, years, remaining_days)) 
```

```python
grader.check("q7")
```

### Question 8
The U.S. Census Bureau projects population based on the following assumptions:

- One birth every 7 seconds
- One death every 13 seconds
- One new immigrant every 45 seconds

Write a program to display the population at the end of the next five years. Assume that the current population is 312, 032, 486, and one year has 365 days. 

```python
SECONDS_IN_A_YEAR = 365 * 24 * 60 * 60  # Number of seconds in one year
SECONDS_IN_5_YEARS = 5 * SECONDS_IN_A_YEAR
BIRTH_RATE = 7  # One birth every 7 seconds
DEATH_RATE = 13  # One death every 13 seconds
IMMIGRATION_RATE = 45  # One immigrant every 45 seconds
    
# Calculate the total number of births, deaths, and immigrants in 5 years
births = ...
deaths = ...
immigrants = ...

current_population = 312032486 

new_population = ...
new_population
```

```python
grader.check("q8")
```

### Question 9

Suppose you save $100 each month into a savings account with an annual interest rate of 5%. Thus, the monthly interest rate is 0.05/12 = 0.00417. After the first month, the value in the account becomes 

$100 * (1 + 0.00417) = 100.417$

After the second month, the value in the account becomes

$(100 + 100.417) * (1 + 0.00417) = 201.252$

After the third month, the value in the account becomes 

$(100 + 201.252) * (1 + 0.00417) = 302.507$
 and so on. 

Write a python program that displays the account value after the sixth month. 

Hint: You will use loop to simply this code in future. 

```python
monthly_deposit = 100  # Amount deposited each month
interest_rate = 0.00417  # Monthly interest rate (5% annual)

# Month-wise account value
first_month = ...
second_month = ...
third_month = ...
fourth_month = ...
fifth_month = ...
sixth_month = ...

# Print the result after the sixth month
print("After the sixth month, the account value is ${:.2f}".format(sixth_month))
```

```python
grader.check("q9")
```

### Question 10
Body Mass Index (BMI) is a measure of health on weight. It can be calculated by taking your weight in kilograms and dividing, by the square of your health in meters. 
Write a python program to calculate a person's BMI given the height is 50 inches, and weight is 95.5. Note, one pound is 0.45359237 kilograms and one inch is 0.0254 meters. 

```python
POUND_TO_KG = 0.45359237  # Conversion factor from pounds to kilograms
INCH_TO_METER = 0.0254  # Conversion factor from inches to meters

weight_pounds = ...
height_inches = ...

# Convert weight and height
weight_kg = ...
height_m = ...

# Calculate BMI
bmi = ...
bmi
```

```python
grader.check("q10")
```

### Question 11
How cold is it outside? The temperature alone is not enough to provide the answer. Other factors including wind speed, relative humidity, and sunshine play important roles in determinig roles in deermining coldness outside. In 2001, the National Weather Service (NWS) implemented the new wind-chill temperator to measure the coldness using temperature and wind speed. The formula is 


$t_{wc} = 35.74 + 0.6215t_a - 35.75v^{0.16} + 0.4275t_av^{0.16}$


where $t_a$ is the outside temperature measured in degrees Fahrenheit, $v$ is the speed measured in miles per hour, and $t_{wc}$ is the wind-chill temperature. The formula cannot be used for wind speeds below 2mph or temperatures below $-58^oF$ or above $41^oF$.

Write a python program that displays the wind-chill temperature given the temperature is $5.3^oF$ and wind speed of 6 mph. 

```python
t_a = 5.3  # Temperature in Fahrenheit
v = 6  # Wind speed in miles per hour

# Wind chill formula
t_wc = ...

# Display the result
print(f"The wind-chill temperature is: {t_wc:.2f}°F")
```

```python
grader.check("q11")
```

### Question 12
If you know the balance and the annual percentage interest rate, you can compute the interest on the next monthly payment using the following formula:

$interest = balance * (annualInterestRate / 1200)$

Write a python program that reads the balance of 1000 and annual percentage interest rate of 3.5 and displays the interest for the next month. 

```python
# Given values
balance = 1000  # Initial balance
annual_interest_rate = 3.5  # Annual interest rate in percentage

interest = ...
interest
```

```python
grader.check("q12")
```

### Question 13
Write a program that takes investment amount of $1000, annual interest rate of 3.25%, number of years of 1 and displays the future investment value using the following formula:

$futureInvestmentValue = investmentAmount * (1 + monthlyInterestRate)^{numberofYears * 12}$

```python
# Given values
investment_amount = 1000  # Initial investment amount
annual_interest_rate = 3.25  # Annual interest rate in percentage
number_of_years = 1  # Number of years

monthly_interest_rate = ...
    
# Calculate future investment value
future_investment_value = ...

# Display the result
print(f"The future investment value is: ${future_investment_value:.2f}")
```

```python
grader.check("q13")
```

### Question 14
Write a python program to calculate the cost of trip. The trip covers a distance of 900.5 miles, and the fuel efficiency of the car in miles per gallon is 25.5, and the price per gallon is $3.55.

```python
# Given values
distance = 900.5  # Distance of the trip in miles
miles_per_gallon = 25.5  # Fuel efficiency in miles per gallon
price_per_gallon = 3.55  # Price per gallon in dollars

gallons_needed = ...

# Calculate the total cost of the trip
total_cost = ...

# Display the result
print(f"The total cost of the trip is: ${total_cost:.2f}")
```

```python
grader.check("q14")
```

### Question 15
Write a program that reads any three-digit integer number between 100 and 9999 and adds all the digits in the integer. For example, if an integer is 932, the sum of all its digits is 14.  

```python
number = int(input("Enter any number between 0 and 1000 in the blank, and hit enter"))
print("You have successfully entered the number {}, now proceed to the next cell".format(number))
```

```python
# Extract the first digit (hundreds place for a three-digit number)
first_digit = ...

# Get the remainder after extracting hundreds place
remaining_digit = ...

# Extract the second and third digits
second_digit = ...
third_digit = ...

# Calculate the sum of digits
sum_of_digits = ...

# Display the result
print(f"The sum of the digits of {number} is: {sum_of_digits}")
```

Tests for q1 (points: None)
```
>>> your_ans == 1024
True
```


Tests for q10 (points: None)
```
>>> math.isclose(bmi, 26.857, rel_tol=0.01) == True
True
```


Tests for q11 (points: None)
```
>>> math.isclose(t_wc, -5.57, rel_tol=0.01)
True
```


Tests for q12 (points: None)
```
>>> math.isclose(interest, 2.92, rel_tol=0.01)
True
```


Tests for q13 (points: None)
```
>>> math.isclose(future_investment_value, 1033, rel_tol=10.0)
True
```


Tests for q14 (points: None)
```
>>> math.isclose(total_cost, 125.36, rel_tol=0.01)
True
```


Tests for q3 (points: None)
```
>>> your_ans == 8.4
True
```


Tests for q4 (points: None)
```
>>> your_ans == 1.83
True
```


Tests for q5 (points: None)
```
>>> your_ans == 22.7
True
```


Tests for q6 (points: None)
```
>>> math.isclose(mark_income, 96800, rel_tol=1e-09) == True
True
```


Tests for q7 (points: None)
```
>>> years == 1902 and remaining_days == 214
True
```


Tests for q8 (points: None)
```
>>> new_population == 325932970
True
```


Tests for q9 (points: None)
```
>>> math.isclose(sixth_month, 608.82, rel_tol=0.01) == True
True
```
