```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab02.ipynb")
```

```python
import numpy as np
import pandas as pd
```

Grading Rubric:

| Question | Points |
| ----- |---- |
| 1.1(a) | 10 |
| 1.1(b) | 10 |
| 1.1(c) | 10 |
| 1.1(d) | 10 |
| 2.1(a) | 10 |
| 2.1(b) | 10 |
| 2.1(c) | 10 |
| 2.1(d) | 10 |
| 2.1(e) | 10 |
| 2.2 | 10 |
| Total | 100 |

# Welcome to Lab02.

Objectives:
1. Intro to Tables
2. Numpy array

### Question 1.1(a) [10 points]

Use `pd.DataFrame` method to create a dataframe, `fruits`, which looks like the following:

<img src="pics/fruits.png" width=200\>


Hint: Use `index` paramter of the `pd.DataFrame` method to add row named labels.

```python
fruits = ...
fruits
```

### Question 1.1(b) [10 points]

Write a Python program (using **`loc`** property) to display the `Fruit Name` and `Price per lbs` for fruits that are typically grown in the summer season. Store your answer in a variable, `summer_fruits` 

Your solution would look like the following:

<img src="pics/summer_fruits.png" width=200/>


```python
summer_fruits = ...
summer_fruits
```

```python
grader.check("q1.1(b)")
```

### Question 1.1(c) [10 points]

Solve the question in 1.1(b) using **`iloc`** property. 

```python
summer_fruits = ...
summer_fruits
```

```python
grader.check("q1.1(c)")
```

### Question 1.1(d) [10 points]

Write a python code to find the most expensive fruit name among the `fruits` dataframe. Store your result in a variable `costly_fruit`. The datatype of the `costly_fruit` is a string. 

```python
costly_fruit = ...
costly_fruit
```

```python
grader.check("q1.1(d)")
```

## 2. Numpy 

<!-- BEGIN QUESTION -->

For each of the Python expressions listed below, write down the output after evaluating the expression. If an error occurs, type **Error**. Provide an explanation or demonstrate your reasoning on how you arrived at the solution, rather than just giving the final answer. Do **NOT** run the code. 

### Question 2.1(a) [10 points]

```
np.array([1, 3, 5, 7]) + np.arange(2, 10, 2)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.1(b) [10 points]

```
np.array([1, 1, 1, 1]) + np.arange(1, 10)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2.1(c) [10 points]

```
np.max(np.arange(15))
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 2.1(d) [10 points]

Give the function `foo`: 

```
def foo(x):
    return x[1] - x[0]
```

What is the output of the following:

```
my_arr = np.array([5, 15, 25, 10])
foo(my_arr)
```

<!-- BEGIN QUESTION -->

### Question 2.1(e) [10 points]

Give the function `foo`: 

```
def foo(x):
    return x[1] - x[0]
```

What is the output of the following:

```
my_arr = np.array([[1, 1, 1], [20, 20, 20], [5, 5, 5]])
foo(my_arr)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 2.2 [10 points]

Create a numpy array, `squares`, containing the first 100 squares i.e, $1^2, 2^2, 3^2, 4^2, 5^2, 6^2, ....., 100^2$  = 1, 4, 9, 16, 25, 36, 49, 64, ...., 10000

```python
squares = ...
squares
```

```python
grader.check("q2.2")
```

Tests for q1.1(b) (points: None)
```
>>> summer_fruits.shape == (2, 2)
True
```
```
>>> len(summer_fruits.columns) == 2
True
```


Tests for q1.1(c) (points: None)
```
>>> summer_fruits.shape == (2, 2)
True
```
```
>>> len(summer_fruits.columns) == 2
True
```


Tests for q1.1(d) (points: None)
```
>>> costly_fruit == 'Apple'
True
```


Tests for q2.2 (points: None)
```
>>> len(squares) == 100
True
```
```
>>> squares[-2].item() == 9801 and squares[-1].item() == 10000
True
```
