```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab05.ipynb")
```

Due Date: 11 Apr 2025 by 11:59 pm

Grading Rubric:

| Question | Points |
| ----- |---- |
| 1(a) | 20 |
| 1(b) | 20 |
| 1(c) | 20 |
| 2 | 20 |
| 3 | 20 |
| Total | 100 |

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
```

Instructions: 

Submission: Please upload the notebook (.ipynb) file on D2L --> Lab -- > Lab05 --> Lab05_Dropbox

1. The lab assignment is individual-based assignment. However, you are encouraged to discuss with your classmates, but DO NOT show your code and copy from others. It will be considered as plagiarism, and proper actions will be taken as per the syllabus.

2. Just because your question has passed the autograder test, doesn't mean necessarily that your answer is correct. So make sure you answer those questions carefully. The instructor has some hidden tests, and only when your question passes those tests successfully during grading, you will be given full points for the question.

3. If you are stuck somewhere, feel free to reach out to me. 

Objectives: To practice and learn about the following topics: 
1. Intro to Probability

# 1. Movie Rental 

Rachel has rented 4 horror, 9 comedy and 2 action movies. Each day, she chooses to watch one movie from
all the movies rented by her. Each movie is equally likely to be chosen by Rachel, and she only watches a
movie once.

Write a mathematical expression (not Python code) to evaluate the probability of the following events happening. You do not need to simplify the arithmetic to find the exact answer, you can keep it in the mathematical
form.

<!-- BEGIN QUESTION -->

### Question 1(a)

What is the probability that Rachel watches comedy movies for the first three days of renting
the movies?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(b)

What is the probability that out of the first two days of renting the movies, she watches
horror one day and action on the other? Note: Here the order in which she watches any of these movies
doesn’t matter.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(c)

What is the probability that Rachel watches at least one horror movie for the first three
days of renting the movie? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2

A certain family has 6 children, consisting of 3 boys and 3 girls. Assuming that all birth orders are equally likely, what is the probability that the 3 eldest children are the girls? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

### Question 3

#### Monty Hall Problem 

Consider the following 7-door version of the Monty Hall problem. There are 7
doors, behind one of which there is a car (which you want), and behind the rest of which
there are goats (which you don’t want). Initially, all possibilities are equally likely for
where the car is. You choose a door. Monty Hall then opens 3 goat doors, and offers
you the option of switching to any of the remaining 3 doors.
Assume that Monty Hall knows which door has the car, will always open 3 goat doors
and offer the option of switching, and that Monty chooses with equal probabilities from
all his choices of which goat doors to open. Should you switch? What is your probability
of success if you switch to one of the remaining 3 doors?


To simulate the scenario, you need to write a Python function called `one_expt` that simulates a single game and returns True if the contestant wins the car by switching, and False otherwise. Then, run this simulation many times (e.g., 50,000 times) to calculate the frequency with which the contestant wins by switching doors.


```python
behind_closed_doors = np.array(['car', 'goat1', 'goat2', 'goat3', 'goat4', 'goat5', 'goat6'])
```

```python
def one_expt():
    np.random.shuffle(behind_closed_doors) # Randomly shuffling the items behind the doors that only Monty knows about

    #contestant_pick is the item that the contestant will pick randomly. Use np.random.choice to select randomly an item from behind closed doors
    contestant_pick = ...

    # this variable stores the option of doors that the contestant will have after monty picks the 3 goat doors randomly. 
    mask = np.logical_and(behind_closed_doors != 'car', behind_closed_doors != contestant_pick)
    doors_left_to_pick_for_monty = behind_closed_doors[mask]

    #monty_pick are the doors that the Monty will pick randomly. Use np.random.choice to select randomly an item from behind closed doors. 
    monty_pick = np.random.choice(doors_left_to_pick_for_monty, ... , replace=...)
    
    doors_left_to_pick_for_contestant = []
    for x in doors_left_to_pick_for_monty:
        if (x not in monty_pick):
            doors_left_to_pick_for_contestant.append(x)

    if(contestant_pick != 'car'):
        doors_left_to_pick_for_contestant.append('car')
            
    contestant_pick_after_switch = ...

    if contestant_pick_after_switch == 'car':
        ...
    else:
        ...
```

```python
repetitions = 50000
total_outcomes  = []

for x in np.arange(repetitions):
    ...
```

```python
total_outcomes = np.array(total_outcomes)
```

```python
prob = ...
prob
```

```python
grader.check("q3")
```

Tests for q3 (points: None)
```
>>> (0.26 < prob < 0.29).item() == True
True
```
