```python
# Initialize Otter
import otter
grader = otter.Notebook("Lab04.ipynb")
```

Due Date: 28 Mar 2025 by 11:59 pm

Grading Rubric:

| Question | Points |
| ----- |---- |
| 1(a) | 10 |
| 1(b) | 10 |
| 1(c) | 10 |
| 1(d) | 10 |
| 1(e) | 10 |
| 1(f) | 10 |
| 2(a) | 10 |
| 2(b) | 10 |
| 2(c) | 20 |
| Total | 100 |

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
```

Instructions: 

Submission: Please upload the notebook (.ipynb) file on D2L --> Lab -- > Lab04 --> Lab04_Dropbox

1. The lab assignment is individual-based assignment. However, you are encouraged to discuss with your classmates, but DO NOT show your code and copy from others. It will be considered as plagiarism, and proper actions will be taken as per the syllabus.

2. Just because your question has passed the autograder test, doesn't mean necessarily that your answer is correct. So make sure you answer those questions carefully. The instructor has some hidden tests, and only when your question passes those tests successfully during grading, you will be given full points for the question.

3. If you are stuck somewhere, feel free to reach out to me. 

Objectives: To practice and learn about the following topics: 
1. Data Visualization
2. Histogram
3. Intro to Probability

# 1. Histogram

Consider the table **nba_salaries** whih contains the NBA players and their salaries for the year 2017-2018. 

```python
nba_salaries = pd.read_csv('data/NBA_season1718_salary.csv', usecols=[1, 2, 3])
nba_salaries.head()
```

### Question 1(a) [10 points]

Create a dataframe **nba** which stores only the two columns: `Player` and `Salary` (converted to millions, rounded upto two decimal points). 

```python
nba = pd.DataFrame({ ... : ...,
                     ... : ...
                   })

nba.head()
```

```python
grader.check("q1a")
```

<!-- BEGIN QUESTION -->

### Question 1(b) [10 points]

Run the following cell to see the salary distribution for the nba players. 

```python
bins = np.array([0, 2, 4, 12, 18, 36])
plt.hist(nba['Salary'], bins = bins, edgecolor = 'black')
plt.xticks(bins)
plt.show()
```

Why is the plot shown above is incorrect?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(c)  [10 points]

Given `outs` and `bins` array in the code below, write python expression to find the percentage of entries in each bins. Store the result in a numpy array `percentage_of_entries`

```python
result = np.histogram(nba['Salary'], bins)
outs = result[0]
bins = result[1]
```

```python
outs
```

```python
bins
```

```python
percentage_of_entries = ...
percentage_of_entries
```

```python
grader.check("q1c")
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(d) [10 points]

What will be the sum of percentages of entries in all bins? Do not write any code for this. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(e) [10 points]

Based on the table given below, which bin contains more players: [2, 4) or [4, 12)? Give reasoning. Do not write any code for this. 

| bin(million dollars) | [0, 2) | [2, 4) | [4, 12) | [12, 18) | [18, 36) |
| --- | --- | --- | --- | --- | --- |
| height (percent per million dollars) | 21.03 | 9.16  | 2.68 | 1.57 | 0.485 | 

Hint: see Question1(c)

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 1(f) [10 points]

To see more detail in the [18, 36) range, the histogram of nba player salaries has been redrawn with bins as shown below. 

| bin(million dollars) | [0, 2) | [2, 4) | [4, 12) | [12, 18) | [18, 26) | [26, 36] |
| --- | --- | --- | --- | --- | --- | --- |
| height (percent per million dollars) | 21.03 | 9.16  | 2.68 | 1.57 | (i) | (ii) |

The expression `nba.shape[0]` evaluates to 573. 

The expression `nba.query('18 <= Salary < 26').shape[0]` evaluates to 37. 

Provide a mathematical expression to find the missing height (i) and (ii) in the above table. Show your solution process. DO NOT write any code for this.

The value of (i) is ______________

The value of (ii) is ________________

```python
nba.shape[0]
```

```python
nba.query('18 <= Salary < 26').shape[0]
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

# 2. Probability 

Please show the solution process and leave the answers in the mathematical form for the following questions.

<!-- BEGIN QUESTION -->

### Question 2(a) [10 points]

Consider a group of four people, how many ways are there to choose a two-person committee?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2(b) [10 points]

Consider a group of four people, how many ways are there to break the people into two teams of two? 

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 2(c) [20 points]

Isaac Newton was consulted about the following problem by Samuel Pepys, who wanted the information for gambling purposes. What is the probability of the following events? 

(i) At least one 6 appears when 6 fair dice are rolled. 

(ii) At least two 6's appear when 12 fair dice are rolled. 

(iii) At least three 6's appear when 18 fair dice are rolled. 

_Type your answer here, replacing this text._

<!-- END QUESTION -->



Tests for q1a (points: None)
```
>>> nba.shape == (573, 2)
True
```
```
>>> len(nba.columns) == 2
True
```
```
>>> np.any('Salary' in nba.columns) == True
np.True_
```


Tests for q1c (points: None)
```
>>> len(percentage_of_entries) == 5
True
```
```
>>> int(percentage_of_entries[-1]) == 8
True
```
