```python
# Initialize Otter
import otter
grader = otter.Notebook("Variables.ipynb")
```

# Variables in Python

<img src="https://smartnumbers.com/wp-content/uploads/2022/01/Blog-Home-Thumbnail-Can-phone-numbers-be-used-to-authenticate-customers-in-the-contact-centre-scaled.jpg"
     alt="Contact"
     style="display: block; margin-left: auto; margin-right: auto; width: 300px;" />
    
    
A python **variable** is a name given to a value. They are used to store data, and are created the moment you assign a value to it. 

For eg: In your phone contacts list, the value is a phone number, and the name of the person you save corresponding to that phone number represents the variable. 

Every variable should have:
1. Name
2. Value
3. Type (Data Types)

## Variable Names

Just like parents follow certain norm to name their child, such as, the names shouldn't be numbers, or shouldn't sound absurd. Similarly, python variable names follows rules given below:
- A variable name must start with a letter or the underscore character
- A variable name cannot start with a number
- A variable name can only contain alpha-numeric characters and underscores (`A-Z`, `a-z`, `0-9`, ` _`)
- Variable names are case-sensitive (name, NaMe, NAME are three different variable names)
- A variable name cannot be any of the [Python keywords](https://www.geeksforgeeks.org/python-keywords/)

Variable names with more than one word can be difficult to read. Therefore, we follow several techniques, such as Camel Case and Snake Case, to make variable names more readable.

#### Camel case
Each word, except the first, starts with a capital letter

#### Snake case
Each word is separated by an underscore character. 

> Use one of the technique to name your multiword variable names. 

For eg: `numberofstudents` is difficult to read. Let's convert it into one of the case

Camel case: `numberOfStudents` 
Snake case: `number_of_students`

## Values

Values represent the data stored in the variable. There are different types of values that can be stored in a variable. Some of those can be classified into the following categories:

| Categories | Example |
| -----      | -----   |
| Text       | "hello how are you", "apple", 'michael jackson' |
| Discrete   | 12395923, -293823593, 5, 0, 239 |
| Continuous | 1238.23595, -297851.5, 5.0, 293.59 |
| Boolean    | True, False |

... etc

### Question 1: 
Create a variable named `personName` and assign your name to it. 

> Every text values in python should always be enclosed with double or single quotes. 

```python
personName = ...
```

```python
grader.check("q1")
```

### Question 2:
Which of the following is a valid variable name? Select all that apply. 

A. _counter <br>
B. numberOfStudent123 <br>
C. DataAnalytics <br>
D. 2023Fall <br>
E. def <br>
F. max_radius <br>
G. #miles <br>
H. data+analytics <br>

Write your options(separated by comma) in the `your_answer` variables. Do not remove the double quotes, or else, your test will fail. 
For eg: If the answers are option A, C, and E then your answer variable will look like `your_answer = "A, C, E"`

```python
your_answer = ...
```

```python
grader.check("q2")
```

## Data Types

Every value (stored in a variable) has an associated data type.

Python has the following in-built data types by default:
| Values | Types |
| ---- | ---- |
| Text | str |
| Discrete | int |
| Continuous | float |
| Boolean | bool |
... etc

You can get the data type of any variable by using the `type()` function

For eg:

```python
x = 5
type(x)
```

```python
y = "hello world"
type(y)
```

```python
z = True
type(z)
```

### Question 3: 
Create a variable named `gpa` and store the value 3.6 into it. 
Use `type()` function to display the type of that variable. 

```python

```

```python
...
```

```python
grader.check("q3")
```

### Why to use variables?

Variables are used to easy access and update the values to be later used in your python code. 

For eg: Suppose you did not save your friend's name in your contact list. In that case, everytime you call your friend, you will have to type all the ten digits. This is a cumbersome and frustrating process. To fix this problem, you save your friend's number with his/her name, so that it is easy to understand and provides you a quick shortcut to dial your friend. Similarly, in order to re-use the data in the variable in later parts of our program, we use variable. 

Moreover, the values stored in the variable can be changed at a later stage of the program, and the variable would store that updated value. Similar to when your friend changes his/her phone number, you update the phone number in your phone, and now calling your friend dials the new updated number. 

<!-- BEGIN QUESTION -->

### Question 4:
What will be the output of the following code? Try to answer it without running the program. 
```
x = 5
x = x + 10
print(x)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

### Question 5: 
What will be the output of the following code? Try to answer it without running the program. 

```
x = 5
y = 2
x = x * 2
y = x + y
print(y)
```

_Type your answer here, replacing this text._

<!-- END QUESTION -->



Tests for q1 (points: None)
```
>>> (type(personName) is str) == True
True
```


Tests for q2 (points: None)
```
>>> ('A' in your_answer and 'B' in your_answer and ('C' in your_answer) and ('F' in your_answer)) == True
True
```


Tests for q3 (points: None)
```
>>> gpa == 3
True
```
```
>>> (type(gpa) is float) == True
True
```
